rule_files:
  - blazegraph.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'blazegraph_free_allocators{instance="wdqs1010:9193"}'
        values: "200000 199990 199980 199970 199960 199950 199940 199930 199920 199910 199900"
      - series: 'blazegraph_free_allocators{instance="wdqs1011:9193"}'
        values: "200000 200000 200000 199999 199999 199999 199998 199998 199998 199997 199996"
      - series: 'blazegraph_free_allocators{instance="unrelated:9194"}'
        values: "200000 199990 199980 199970 199960 199950 199940 199930 199920 199910 199900"
    alert_rule_test:
      - alertname: BlazegraphFreeAllocatorsDecreasingRapidly
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              instance: "wdqs1010:9193"
            exp_annotations:
              summary: "Blazegraph instance wdqs1010:9193 is burning free allocators at a very high rate"
              description: "Blazegraph is misbehaving and will rapidly corrupt its journal."
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Runbook#Free_allocators_decrease_rapidly
              dashboard: https://grafana.wikimedia.org/d/000000489/wikidata-query-service

  - interval: 1m
    input_series:
      - series: 'blazegraph_free_allocators{instance="wdqs1001:9195"}'
        values: '95000+0x10'
    alert_rule_test:
      - alertname: BlazegraphFreeAllocatorsRunningOut
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              instance: "wdqs1001:9195"
            exp_annotations:
              summary: "Blazegraph instance wdqs1001:9195 is running out of free allocators"
              description: "Blazegraph will soon be unable to apply updates"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Runbook
              dashboard: https://grafana.wikimedia.org/d/000000489/wikidata-query-service

  - interval: 1m
    external_labels:
      site: eqiad
    input_series:
      - series: 'blazegraph_lastupdated{site="eqiad", cluster="wdqs-main", instance="wdqs3050:9193"}'
        values: "0x10"
      - series: org_wikidata_query_rdf_blazegraph_filters_QueryEventSenderFilter_event_sender_filter_StartedQueries{instance="wdqs3050:9102", cluster="wdqs-main", site="eqiad"}
        values: "10000+10000x10"
    alert_rule_test:
      - alertname: ElevatedMaxLagWDQS
        eval_time: 12m
        exp_alerts:
          - exp_labels:
              team: search-platform
              severity: critical
            exp_annotations:
              summary: "WDQS lag is above 10 minutes"
              description: "WDQS cluster in eqiad is throttling all connections due to high lag"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Runbook
              dashboard: https://grafana.wikimedia.org/d/000000489/wikidata-query-service
