rule_files:
  - cirrussearch_k8s.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: mediawiki_CirrusSearch_pool_counter_seconds_count{status="success", type="fulltext"}
        values: '0+120000x6'
      - series: mediawiki_CirrusSearch_pool_counter_seconds_count{status="failure", type="fulltext"}
        values: '0+30001x6'
      - series: mediawiki_CirrusSearch_pool_counter_seconds_count{status="failure", type="completion"}
        values: '0+30000x6'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusSearchPoolCounterRejectionTooHigh
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
            exp_annotations:
              summary: 'MediaWiki CirrusSearch failing to obtain a token from the pool counter at a very high rate'
              description: "Rate of failures when obtaining a token from the pool counter is above 1000 errors/sec"
              runbook: 'https://wikitech.wikimedia.org/wiki/Search/Elasticsearch_Administration#Pool_Counter_rejections_(search_is_currently_too_busy)'
              dashboard: 'https://grafana.wikimedia.org/d/qrOStmdGk/elasticsearch-pool-counters?viewPanel=4&orgId=1'
  - interval: '1d'
    input_series:
      - series: mediawiki_CirrusSearch_sanitization_total{action="fixed", search_cluster="broken_cluster"}
        values: '0+20000x7'
      - series: mediawiki_CirrusSearch_sanitization_total{action="fixed", search_cluster="sane_cluster"}
        values: '0+20x7'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusSearchSaneitizerFixRateTooHigh
        eval_time: '7d'
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              search_cluster: broken_cluster
            exp_annotations:
              summary: 'MediaWiki CirrusSearch Saneitizer is fixing an abnormally high number of documents in broken_cluster'
              description: "MediaWiki CirrusSearch Saneitizer Weekly Fix Rate is above 100000/week for broken_cluster"
              runbook: 'https://wikitech.wikimedia.org/wiki/Search/CirrusStreamingUpdater#San(e)itizing'
              dashboard: 'https://grafana.wikimedia.org/d/JLK3I_siz/elasticsearch-indexing?viewPanel=59&orgId=1&from=now-6M&to=now&var-search_cluster=broken_cluster'
  - interval: '1m'
    input_series:
      # histograms for full_text search to eqiad p95 is 1.238s
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="full_text", le="+Inf"}
        values: '0+1x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="full_text", le="2.5"}
        values: '0+101x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="full_text", le="1"}
        values: '0+95x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="full_text", le="0.5"}
        values: '0+10x6'

      # histograms for full_text search to codfw p95 is 972.3ms
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="full_text", le="+Inf"}
        values: '0+1x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="full_text", le="2.5"}
        values: '0+95x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="full_text", le="1"}
        values: '0+101x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="full_text", le="0.5"}
        values: '0+10x6'

      # histograms for comp_suggest search to eqiad p95 is 289.6ms
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="comp_suggest", le="+Inf"}
        values: '0+1x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="comp_suggest", le="0.5"}
        values: '0+101x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="comp_suggest", le="0.25"}
        values: '0+95x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="comp_suggest", le="0.1"}
        values: '0+10x6'

      # histograms for comp_suggest search to codfw p95 is 241.7ms
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="comp_suggest", le="+Inf"}
        values: '0+1x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="comp_suggest", le="0.5"}
        values: '0+95x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="comp_suggest", le="0.25"}
        values: '0+101x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="comp_suggest", le="0.1"}
        values: '0+10x6'

      # histograms for more_like search to eqiad p95 is 1.75s
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="more_like", le="+Inf"}
        values: '0+1x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="more_like", le="2.5"}
        values: '0+100x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="eqiad", type="more_like", le="1"}
        values: '0+90x6'

      # histograms for more_like search to codfw p95 is 1.429s
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="more_like", le="+Inf"}
        values: '0+1x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="more_like", le="2.5"}
        values: '0+100x6'
      - series: mediawiki_CirrusSearch_request_time_seconds_bucket{search_cluster="codfw", type="more_like", le="1"}
        values: '0+93x6'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusSearchFullTextLatencyTooHigh
        eval_time: '6m'
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              search_cluster: eqiad
            exp_annotations:
              summary: 'CirrusSearch full_text 95th percentiles latency is too high (mw@eqiad to eqiad)'
              description: "CirrusSearch full_text eqiad 95th percentile latency"
              runbook: 'https://wikitech.wikimedia.org/wiki/Search#Health/Activity_Monitoring'
              dashboard: 'https://grafana.wikimedia.org/d/dc04b9f2-b8d5-4ab6-9482-5d9a75728951/elasticsearch-percentiles?orgId=1'
      - alertname: CirrusSearchCompletionLatencyTooHigh
        eval_time: '6m'
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              search_cluster: eqiad
            exp_annotations:
              summary: 'CirrusSearch comp_suggest 95th percentiles latency is too high (mw@eqiad to eqiad)'
              description: "CirrusSearch comp_suggest eqiad 95th percentile latency"
              runbook: 'https://wikitech.wikimedia.org/wiki/Search#Health/Activity_Monitoring'
              dashboard: 'https://grafana.wikimedia.org/d/dc04b9f2-b8d5-4ab6-9482-5d9a75728951/elasticsearch-percentiles?orgId=1'
      - alertname: CirrusSearchMoreLikeLatencyTooHigh
        eval_time: '6m'
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              search_cluster: eqiad
            exp_annotations:
              summary: 'CirrusSearch more_like 95th percentiles latency is too high (mw@eqiad to eqiad)'
              description: "CirrusSearch more_like eqiad 95th percentile latency"
              runbook: 'https://wikitech.wikimedia.org/wiki/Search#Health/Activity_Monitoring'
              dashboard: 'https://grafana.wikimedia.org/d/dc04b9f2-b8d5-4ab6-9482-5d9a75728951/elasticsearch-percentiles?orgId=1'
  - interval: '1m'
    input_series:
      - series: mediawiki_CirrusSearch_backend_failures_total{type="memory_issue", search_cluster="eqiad"}
        values: '0+3x12'
      - series: mediawiki_CirrusSearch_backend_failures_total{type="memory_issue", search_cluster="codfw"}
        values: '0+2x12'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusSearchBackendMemoryIssue
        eval_time: '12m'
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              search_cluster: eqiad
            exp_annotations:
              summary: 'CirrusSearch backend failed 30 times in the last 10 minutes due to memory usage (mw@eqiad to eqiad)'
              description: 'Number of requests triggering circuit breakers due to excessive memory usage'
              runbook: 'https://wikitech.wikimedia.org/wiki/Search#Health/Activity_Monitoring'
              dashboard: 'https://grafana.wikimedia.org/d/dc04b9f2-b8d5-4ab6-9482-5d9a75728951/elasticsearch-percentiles?orgId=1'

  - interval: 1m
    input_series:
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="search-https_eqiad", app="mediawiki", envoy_response_code="503"}
        values: '0+400x20'
    external_labels:
      site: eqiad
      prometheus: ops
    alert_rule_test:
      - alertname: CirrusBackendErrorRateTooHigh
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
            exp_annotations:
                  summary: 'CirrusSearch getting over 0.5% error responses from elasticsearch'
                  description: "Envoy reports 5xx rate over 5 minutes of more than .5%"
                  runbook: 'TODO'
                  dashboard: 'https://grafana.wikimedia.org/d/b1jttnFMz/envoy-telemetry-k8s?orgId=1&var-datasource=thanos&var-site=eqiad&var-prometheus=k8s&var-app=mediawiki&var-destination=search-chi-eqiad&var-destination=search-omega-eqiad&var-destination=search-psi-eqiad&var-destination=search-chi-codfw&var-destination=search-omega-codfw&var-destination=search-psi-codfw'
