rule_files:
  - cirrus_streaming_updater.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_job_uptime{job_name="cirrus_streaming_updater_producer_eqiad", kubernetes_namespace="cirrus-streaming-updater", pod="1", release="producer"}'
        values: "2000000 0 60000 120000 180000 0 0 0 0 0 0 0"
      - series: 'flink_jobmanager_job_uptime{job_name="cirrus_streaming_updater_producer_eqiad", kubernetes_namespace="cirrus-streaming-updater", pod="2", release="producer"}'
        values: "0 0 0 0 0 0 60000 120000 180000 240000 0 60000"
      - series: 'flink_taskmanager_job_task_operator_processing_latency_ms{kubernetes_namespace="cirrus-streaming-updater", job_name="cirrus_streaming_updater_producer_eqiad", release="producer"}'
        values: "600001"
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: CirrusStreamingUpdaterFlinkJobUnstable
        eval_time: 12m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: cirrus-streaming-updater
              job_name: cirrus_streaming_updater_producer_eqiad
              release: producer
            exp_annotations:
              summary: "cirrus_streaming_updater_producer_eqiad in eqiad (k8s) is unstable"
              description: "The flink job has an uptime below 5 minutes for the past 10 minutes, it is probably trying to start in vain"
              runbook: "https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater"
              dashboard: https://grafana.wikimedia.org/d/K9x0c4aVk/flink-app?var-datasource=eqiad+prometheus%2Fk8s&var-namespace=cirrus-streaming-updater&var-helm_release=producer
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_job_uptime{job_name="Unrelated_Job", kubernetes_namespace="unrelated-ns", release="producer"}'
        values: "2000000"
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: CirrusProducerFlinkJobNotRunning
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: cirrus-streaming-updater
            exp_annotations:
              summary: "cirrus_streaming_updater_producer in eqiad (k8s) is not running"
              description: "The flink job is not running"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/K9x0c4aVk/flink-app?var-datasource=eqiad+prometheus%2Fk8s&var-namespace=cirrus-streaming-updater&var-helm_release=producer
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_numRegisteredTaskManagers{job_name="cirrus-streaming-updater-producer", kubernetes_namespace="cirrus-streaming-updater", release="producer"}'
        values: "0x6"
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: CirrusStreamingUpdaterFlinkNoRegisteredTask
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: cirrus-streaming-updater
              release: producer
              job_name: cirrus-streaming-updater-producer
            exp_annotations:
              summary: "cirrus-streaming-updater job cirrus-streaming-updater-producer in eqiad (k8s) is running without any taskmanagers"
              description: "The flink job is running but without any taskmanagers (probably doing nothing)"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/K9x0c4aVk/flink-app?var-datasource=eqiad+prometheus%2Fk8s&var-namespace=cirrus-streaming-updater&var-helm_release=producer
  - interval: 1m
    input_series:
      - series: 'flink_taskmanager_job_task_operator_numRecordsOut{job_name="cirrus_streaming_updater_consumer_search_eqiad", kubernetes_namespace="cirrus-streaming-updater", operator_name="fetch_failure_sink:_Writer", release="consumer-search"}'
        values: '1+1x21'
      - series: 'flink_taskmanager_job_task_operator_numRecordsOut{job_name="cirrus_streaming_updater_consumer_cloudelastic_search_eqiad", kubernetes_namespace="cirrus-streaming-updater", operator_name="fetch_failure_sink:_Writer", release="consumer-cloudelastic"}'
        values: '1+30x21'
      - series: 'flink_taskmanager_job_task_operator_numRecordsOut{job_name="cirrus_streaming_updater_consumer_backfill_search_eqiad", kubernetes_namespace="cirrus-streaming-updater", operator_name="fetch_failure_sink:_Writer", release="consumer-search-backfill"}'
        values: '1+30x21'
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: CirrusConsumerFetchErrorRate
        eval_time: 21m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: cirrus-streaming-updater
              job_name: cirrus_streaming_updater_consumer_cloudelastic_search_eqiad
              release: consumer-cloudelastic
            exp_annotations:
              summary: "cirrus_streaming_updater_consumer_cloudelastic_search_eqiad in eqiad (k8s): fetch error rate too high"
              description: "The flink job is failing to fetch the cirrus documents at a high rate"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/K9x0c4aVk/flink-app?var-datasource=eqiad+prometheus%2Fk8s&var-namespace=cirrus-streaming-updater&var-helm_release=consumer-cloudelastic

  - interval: 1m
    input_series:
      - series: 'flink_taskmanager_job_task_operator_numRecordsOut{job_name="cirrus_streaming_updater_consumer_search_eqiad", kubernetes_namespace="cirrus-streaming-updater", operator_name="rerender_fetch_failure_sink:_Writer", release="consumer-search"}'
        values: '1+1x21'
      - series: 'flink_taskmanager_job_task_operator_numRecordsOut{job_name="cirrus_streaming_updater_consumer_cloudelastic_search_eqiad", kubernetes_namespace="cirrus-streaming-updater", operator_name="rerender_fetch_failure_sink:_Writer", release="consumer-cloudelastic"}'
        values: '1+100x21'
      - series: 'flink_taskmanager_job_task_operator_numRecordsOut{job_name="cirrus_streaming_updater_consumer_backfill_search_eqiad", kubernetes_namespace="cirrus-streaming-updater", operator_name="rerender_fetch_failure_sink:_Writer", release="consumer-search-backfill"}'
        values: '1+100x21'
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: CirrusConsumerRerenderFetchErrorRate
        eval_time: 21m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: cirrus-streaming-updater
              job_name: cirrus_streaming_updater_consumer_cloudelastic_search_eqiad
              release: consumer-cloudelastic
            exp_annotations:
              summary: "cirrus_streaming_updater_consumer_cloudelastic_search_eqiad in eqiad (k8s): fetch error (rerenders) rate too high"
              description: "The flink job is failing to fetch the cirrus documents at a high rate for rerenders"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/K9x0c4aVk/flink-app?var-datasource=eqiad+prometheus%2Fk8s&var-namespace=cirrus-streaming-updater&var-helm_release=consumer-cloudelastic

  - interval: 1m
    input_series:
      - series: 'flink_taskmanager_job_task_operator_bulk_change_type_result_actions{app="flink-app-consumer-cloudelastic"}'
        values: 0 0 0 0 0 0 0 0 0 0 0 0
    alert_rule_test:
      - alertname: CirrusStreamingUpdaterRateTooLow
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              app: flink-app-consumer-cloudelastic
            exp_annotations:
              summary: "CirrusSearch update rate from flink-app-consumer-cloudelastic is critically low"
              description: "Updates are either not flowing, or flowing at a greatly reduced rate."
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/jKqki4MSk/cirrus-streaming-updater

  - interval: 1m
    input_series:
      - series: 'flink_taskmanager_job_task_operator_weighted_tags_tag_prefix_set{job_name="cirrus_streaming_updater_consumer_search_eqiad", release="consumer-search"}'
        values: '0+0x11'
      - series: 'flink_taskmanager_job_task_operator_weighted_tags_tag_prefix_set{job_name="cirrus_streaming_updater_consumer_cloudelastic_eqiad", release="consumer-cloudelastic"}'
        values: '1+10x11'
      - series: 'flink_taskmanager_job_task_operator_weighted_tags_tag_prefix_set{job_name="cirrus_streaming_updater_consumer_search_eqiad_backfill", release="consumer-search-backfill"}'
        values: '0+0x11'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusStreamingUpdaterSetWeightedTagsTooLow
        eval_time: 11m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              release: consumer-search
            exp_annotations:
              summary: "CirrusSearch consumer-search@eqiad is setting too few weighted tags"
              description: "The number of weighted_tags set is too low in the last 10 minutes"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/fe251f4f-f6cf-4010-8d78-5f482255b16f/cirrussearch-update-pipeline-weighted-tags?orgId=1&var-tag_prefix=All&var-search_cluster_site=eqiad&var-search_cluster=consumer-search
  - interval: 1m
    input_series:
      - series: 'flink_taskmanager_job_task_operator_weighted_tags_tag_prefix_clear{job_name="cirrus_streaming_updater_consumer_search_eqiad", release="consumer-search"}'
        values: '0+0x11'
      - series: 'flink_taskmanager_job_task_operator_weighted_tags_tag_prefix_clear{job_name="cirrus_streaming_updater_consumer_cloudelastic_eqiad", release="consumer-cloudelastic"}'
        values: '1+10x11'
      - series: 'flink_taskmanager_job_task_operator_weighted_tags_tag_prefix_clear{job_name="cirrus_streaming_updater_consumer_search_eqiad_backfill", release="consumer-search-backfill"}'
        values: '0+0x11'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusStreamingUpdaterClearWeightedTagsTooLow
        eval_time: 11m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              release: consumer-search
            exp_annotations:
              summary: "CirrusSearch consumer-search@eqiad is clearing too few weighted tags"
              description: "The number of weighted_tags cleared is too low in the last 10 minutes"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/fe251f4f-f6cf-4010-8d78-5f482255b16f/cirrussearch-update-pipeline-weighted-tags?orgId=1&var-tag_prefix=All&var-search_cluster_site=eqiad&var-search_cluster=consumer-search
