rule_files:
  - cirrus_streaming_updater_cloudelastic.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_job_uptime{job_name="Unrelated_Job", kubernetes_namespace="unrelated-ns", release="consumer-cloudelastic"}'
        values: "2000000"
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: CirrusConsumerCloudelasticFlinkJobNotRunning
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: cirrus-streaming-updater
            exp_annotations:
              summary: "cirrus_streaming_updater_cloudelastic_consumer in eqiad (k8s) is not running"
              description: "The flink job is not running"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/K9x0c4aVk/flink-app?var-datasource=eqiad+prometheus%2Fk8s&var-namespace=cirrus-streaming-updater&var-helm_release=consumer-cloudelastic
