rule_files:
  - cirrussearch.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'elasticsearch_jvm_memory_pool_used_bytes{exported_cluster="production-search-eqiad", pool="young", name="elastic1001-production-search-eqiad"}'
        values: 1000000000+0x10
      - series: 'elasticsearch_jvm_memory_pool_used_bytes{exported_cluster="production-search-eqiad", pool="young", name="elastic1002-production-search-eqiad"}'
        values: 0+0x10
    alert_rule_test:
      - alertname: CirrusSearchJVMGCYoungPoolInsufficient
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              exported_cluster: production-search-eqiad
              pool: young
              name: elastic1002-production-search-eqiad
            exp_annotations:
              summary: "Elasticsearch instance elastic1002-production-search-eqiad is showing memory pressure in the young pool"
              description: "Elasticsearch is out of memory and will have degraded latency, reject requests, and could OOM"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Stuck_in_old_GC_hell
              dashboard: https://grafana.wikimedia.org/d/000000462/elasticsearch-memory

  - interval: 1m
    input_series:
      - series: elasticsearch_jvm_gc_collection_seconds_count{exported_cluster="production-search-testing", gc="old", name="elastic1002-production-search-eqiad"}
        values: 0+2x60
    alert_rule_test:
      - alertname: CirrusSearchHighOldGCFrequency
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              name: elastic1002-production-search-eqiad
              gc: old
              exported_cluster: production-search-testing
            exp_annotations:
              summary: "Elasticsearch instance elastic1002-production-search-eqiad is running the old gc excessively"
              description: "Elasticsearch is out of memory and will have degraded latency"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Stuck_in_old_GC_hell
              dashboard: https://grafana.wikimedia.org/d/000000462/elasticsearch-memory

  - interval: 1m
    input_series:
     - series: elasticsearch_indices_age_days{index="commonswiki_file", exported_cluster="alertmanager-test"}
       values: 121
    alert_rule_test:
     - alertname: CirrusSearchIndexTooOld
       eval_time: 1m
       exp_alerts:
        - exp_labels:
            severity: task
            team: search-platform
            index: commonswiki_file
            exported_cluster: alertmanager-test
          exp_annotations:
            summary: "commonswiki_file index in alertmanager-test is ready for reindexing"
            description: "The reindexing process should be run to ensure multi-lingual analysis chain updates are applied"
            runbook: TODO
            dashboard: TODO

  - interval: 1m
    input_series:
      - series: kafka_burrow_partition_lag{exported_cluster="main-eqiad",group="cpjobqueue-low_traffic_jobs", topic="eqiad.mediawiki.job.cirrusSearchOtherIndex",partition="0"}
        values: '210000+0x10'
      - series: kafka_burrow_partition_lag{exported_cluster="main-eqiad",group="cpjobqueue-cirrusSearchLinksUpdate", topic="eqiad.mediawiki.job.cirrusSearchLinksUpdate",partition="0"}
        values: 21000
      - series: kafka_burrow_partition_lag{exported_cluster="main-eqiad",group="cpjobqueue-cirrusSearchElasticaWrite_partitioner", topic="eqiad.mediawiki.job.cirrusSearchElasticaWrite",partition="0"}
        values: 220000
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CirrusSearchJobQueueBacklogTooBig
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              group: cpjobqueue-low_traffic_jobs
              partition: 0
              topic: eqiad.mediawiki.job.cirrusSearchOtherIndex
              exported_cluster: main-eqiad
              site: eqiad
            exp_annotations:
              summary: "CirrusSearch job topic eqiad.mediawiki.job.cirrusSearchOtherIndex is heavily backlogged with 210k messages"
              description: "CirrusSearch job backlog is too big"
              runbook: TODO
              dashboard: "https://grafana.wikimedia.org/d/CbmStnlGk/jobqueue-job?orgId=1&var-dc=eqiad%20prometheus/k8s&var-job=cirrusSearchOtherIndex"

  - interval: 1m
    input_series:
      - series: kafka_server_BrokerTopicMetrics_MessagesIn_total{topic="eqiad.cirrussearch.update_pipeline.update.rc0"}
        values: '0x11'
    external_labels:
      site: eqiad
      prometheus: ops
    alert_rule_test:
      - alertname: CirrusSearchUpdaterKafkaMessagesInTooLow
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
            exp_annotations:
                  summary: 'The summed message update rate of `(eqiad|codfw).cirrussearch.update_pipeline.update.rc0` is too low'
                  description: "Unexpectedly low Kafka topic message-in rate for the last 5 minutes"
                  runbook: 'TODO'
                  dashboard: 'https://grafana.wikimedia.org/d/000000234/kafka-by-topic?orgId=1&var-dc=eqiad%2520prometheus%252Fops&var-kafka_cluster=main-eqiad&var-kafka_broker=All&from=now-1h&to=now&refresh=5m&var-topic=codfw.cirrussearch.update_pipeline.update.rc0&var-topic=eqcodfw.cirrussearch.update_pipeline.update.rc0&var-topic=eqiad.cirrussearch.update_pipeline.update.rc0&viewPanel=6'


  - interval: 1m
    input_series:
      - series: elasticsearch_indices_indexing_index_total{exported_cluster="production-search-eqiad", index="", name="elastic1337-production-search-eqiad"}
        values: 0+0x30
    alert_rule_test:
      - alertname: CirrusSearchNodeIndexingNotIncreasing
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              name: elastic1337-production-search-eqiad
              exported_cluster: production-search-eqiad
            exp_annotations:
              summary: "Elasticsearch instance elastic1337-production-search-eqiad is not indexing"
              description: "An Elasticsearch instance is not performing index work; find out why"
              runbook: https://wikitech.wikimedia.org/wiki/Search#Indexing_hung_and_not_making_progress
              dashboard: 'https://grafana.wikimedia.org/d/JLK3I_siz/elasticsearch-indexing?orgId=1&from=now-3d&to=now&viewPanel=57'
  - interval: 1m
    input_series:
      - series: elasticsearch_indices_age_days{exported_cluster="production-search-psi-codfw", index="*_titlesuggest"}
        values: 5+0x10
    alert_rule_test:
      - alertname: CirrusSearchTitleSuggestIndexTooOld
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              index: '*_titlesuggest'
              exported_cluster: production-search-psi-codfw
            exp_annotations:
              summary: "Some search indices that power autocomplete have not been updated recently"
              description: "The daily build process for autocomplete indices is probably having issues and should be investigated."
              runbook: https://wikitech.wikimedia.org/wiki/Search#CirrusSearch_titlesuggest_index_is_too_old
              dashboard: 'TODO'

  - interval: 1m
    input_series:
      - series: envoy_cluster_upstream_rq{envoy_cluster_name="search-https_eqiad", cluster="appserver", envoy_response_code="503"}
        values: '0+400x20'
    external_labels:
      site: eqiad
      prometheus: ops
    alert_rule_test:
      - alertname: CirrusBackendErrorRateTooHigh
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
            exp_annotations:
                  summary: 'CirrusSearch getting over 0.5% error responses from elasticsearch'
                  description: "Envoy reports 5xx rate over 5 minutes of more than .5%"
                  runbook: 'TODO'
                  dashboard: 'https://grafana.wikimedia.org/d/VTCkm29Wz/envoy-telemetry?orgId=1&var-datasource=eqiad%20prometheus-ops&var-origin=appserver&var-origin_instance=All&var-destination=search-https_eqiad'
