# https://prometheus.io/docs/prometheus/latest/configuration/unit_testing_rules/
# Test locally using `promtool test rules rdf-streaming-updater-test.yaml`
rule_files:
  - rdf_streaming_updater.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_job_uptime{job_name="WDQS_Streaming_Updater", kubernetes_namespace="rdf-streaming-updater", pod="1"}'
        values: "2000000 0 60000 120000 180000 0 0 0 0 0 0 0"
      - series: 'flink_jobmanager_job_uptime{job_name="WDQS_Streaming_Updater", kubernetes_namespace="rdf-streaming-updater", pod="2"}'
        values: "0 0 0 0 0 0 60000 120000 180000 240000 0 60000"
      - series: 'flink_taskmanager_job_task_operator_processing_latency_ms{kubernetes_namespace="rdf-streaming-updater", job_name="WCQS_Streaming_Updater"}'
        values: "600001"
    external_labels:
      site: codfw
      prometheus: k8s
    alert_rule_test:
      - alertname: RdfStreamingUpdaterFlinkJobUnstable
        eval_time: 12m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: rdf-streaming-updater
              job_name: "WDQS_Streaming_Updater"
            exp_annotations:
              summary: "WDQS_Streaming_Updater in codfw (k8s) is unstable"
              description: "The flink job has an uptime below 5 minutes for the past 10 minutes, it is probably trying to start in vain"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/gCFgfpG7k/flink-session-cluster
      - alertname: RdfStreamingUpdaterFlinkProcessingLatencyIsHigh
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              kubernetes_namespace: rdf-streaming-updater
              job_name: "WCQS_Streaming_Updater"
            exp_annotations:
              summary: "Processing latency of WCQS_Streaming_Updater in codfw (k8s) is above 10 minutes"
              description: "The flink job has a processing latency above 10 minutes"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/fdU5Zx-Mk/wdqs-streaming-updater
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_taskSlotsTotal{kubernetes_namespace="rdf-streaming-updater"}'
        values: 22
      - series: 'flink_jobmanager_taskSlotsTotal{kubernetes_namespace="unrelated-session-cluster"}'
        values: 4
    external_labels:
      site: codfw
      prometheus: k8s
    alert_rule_test:
      - alertname: RdfStreamingUpdaterNotEnoughTaskSlots
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: search-platform
              kubernetes_namespace: rdf-streaming-updater
            exp_annotations:
              summary: "The flink session cluster rdf-streaming-updater in codfw (k8s) does not have enough task slots"
              description: "The expected number of task slots is too low for the jobs to function properly"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/gCFgfpG7k/flink-session-cluster
  - interval: 1m
    input_series:
      - series: 'flink_jobmanager_job_uptime{job_name="Unrelated_Job", kubernetes_namespace="unrelated-ns"}'
        values: "2000000"
    external_labels:
      site: eqiad
      prometheus: k8s
    alert_rule_test:
      - alertname: WdqsStreamingUpdaterFlinkJobNotRunning
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: rdf-streaming-updater
              job_name: "WDQS_Streaming_Updater"
            exp_annotations:
              summary: "WDQS_Streaming_Updater in eqiad (k8s) is not running"
              description: "The flink job is not running"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/gCFgfpG7k/flink-session-cluster
      - alertname: WcqsStreamingUpdaterFlinkJobNotRunning
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: search-platform
              kubernetes_namespace: rdf-streaming-updater
              job_name: "WCQS_Streaming_Updater"
            exp_annotations:
              summary: "WCQS_Streaming_Updater in eqiad (k8s) is not running"
              description: "The flink job is not running"
              runbook: https://wikitech.wikimedia.org/wiki/Wikidata_Query_Service/Streaming_Updater
              dashboard: https://grafana.wikimedia.org/d/gCFgfpG7k/flink-session-cluster
