rule_files:
  - airflow-k8s.yaml
tests:
  - input_series:
    - series: kube_deployment_status_replicas_available{deployment="airflow-webserver", namespace="airflow-test-k8s"}
      values: 0
    - series: kube_deployment_status_replicas_available{deployment="airflow-scheduler", namespace="airflow-test-k8s"}
      values: 1
    - series: kube_deployment_status_replicas_available{deployment="airflow-kerberos", namespace="airflow-test-k8s"}
      values: 1
    alert_rule_test:
    - alertname: AirflowDeploymentUnavailable
      eval_time: 5m
      exp_alerts:
      - exp_annotations:
          description: airflow-webserver.airflow-test-k8s is unavailable on dse-k8s-eqiad
          summary: airflow-webserver.airflow-test-k8s is unavailable on dse-k8s-eqiad
          dashboard: "https://grafana.wikimedia.org/d/hyl18XgMk/kubernetes-container-details?orgId=1&var-datasource=eqiad%2Bprometheus/k8s-dse&var-namespace=airflow-test-k8s&var-container=All"
          runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Airflow/Kubernetes#Pods_are_not_running"
        exp_labels:
          namespace: airflow-test-k8s
          deployment: airflow-webserver
          severity: critical
          team: data-platform
  - input_series:
    - series: airflow_dagbag_size{namespace="airflow-test-k8s"}
      values: 0
    alert_rule_test:
    - alertname: AirflowDagBagEmpty
      eval_time: 5m
      exp_alerts:
      - exp_annotations:
          description: DagBag is empty on airflow-test-k8s.dse-k8s-eqiad
          summary: DagBag is empty on airflow-test-k8s.dse-k8s-eqiad
          dashboard: "https://grafana.wikimedia.org/d/lFXqBGxWk/airflow-instance?&var-airflow_instance=airflow-test-k8s"
          runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Airflow/Kubernetes#An_instance_DAG_Bag_is_empty"
        exp_labels:
          namespace: airflow-test-k8s
          severity: critical
          team: data-platform
