rule_files:
  - cephosd_cluster.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: ceph_health_status{job="ceph"}
        values: "0 1 1 2+0x7"
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CephOSDClusterInError
        eval_time: 8m
        exp_alerts:
          - exp_labels:
              team: data-platform
              severity: critical
              job: ceph
            exp_annotations:
              description: "eqiad Cephosd cluster is unable to serve traffic."
              summary: "#warning Cephosd cluster in eqiad is in error status"
              dashboard: "https://grafana.wikimedia.org/d/tbO9LAiZK/cephosd-cluster?orgId=1&search=open&tag=cephosd"
              runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Ceph"

  - interval: 1m
    input_series:
      - series: ceph_health_status{job="ceph"}
        values: "0 1+0x9"
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CephOSDClusterInWarning
        eval_time: 7m
        exp_alerts:
          - exp_labels:
              team: data-platform
              severity: warning
              job: ceph
            exp_annotations:
              description: "High availability of the Ceph cluster in eqiad is compromised, but things should still be working as expected."
              summary: "Ceph cluster in eqiad is in warning status"
              dashboard: "https://grafana.wikimedia.org/d/tbO9LAiZK/cephosd-cluster?orgId=1&search=open&tag=cephosd"
              runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Ceph"

  - interval: 1m
    input_series: []
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CephOSDClusterInUnknown
        eval_time: 8m
        exp_alerts:
          - exp_labels:
              team: data-platform
              severity: critical
            exp_annotations:
              description: "Unable to gather metrics from eqiad Ceph cluster."
              summary: "#warning Ceph cluster in eqiad is in unknown status"
              dashboard: "https://grafana.wikimedia.org/d/tbO9LAiZK/cephosd-cluster?orgId=1&search=open&tag=cephosd"
              runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Ceph"

  - interval: 1m
    input_series:
      - series: ceph_mon_quorum_status{job="ceph", ceph_daemon="mon.cephosd1001"}
        values: "1 1 0x10"
      - series: ceph_mon_quorum_status{job="ceph", ceph_daemon="mon.cephosd1002"}
        values: "1 0x10"
      - series: ceph_mon_quorum_status{job="ceph", ceph_daemon="mon.cephosd1003"}
        values: "1 1 1 0x10"
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CephOSDClusterAllMonsDown
        eval_time: 8m
        exp_alerts:
          - exp_labels:
              team: data-platform
              severity: critical
            exp_annotations:
              description: "All monitor nodes for the ceph cluster in eqiad are down (quorum lost), the cluster is down"
              summary: "#warning All the monitor nodes for the cluster in eqiad have lost quorum"
              dashboard: "https://grafana.wikimedia.org/d/tbO9LAiZK/cephosd-cluster?orgId=1&search=open&tag=cephosd"
              runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Ceph"

  - interval: 1m
    input_series:
      - series: ceph_healthcheck_slow_ops{job="ceph"}
        values: "0 1 2 3 4 5"
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: CephOSDSlowOps
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              team: data-platform
              severity: critical
              job: ceph
            exp_annotations:
              description: "Ceph cluster in eqiad has slow ops, which might be blocking some writes"
              summary: "Ceph cluster in eqiad has 5 slow ops"
              dashboard: "https://grafana.wikimedia.org/d/tbO9LAiZK/cephosd-cluster?orgId=1&search=open&tag=cephosd"
              runbook: "https://wikitech.wikimedia.org/wiki/Data_Platform/Systems/Ceph"
