rule_files:
  - datahub-availability.yaml
tests:
  - input_series:
    - series: kube_deployment_status_replicas_available{deployment="datahub-gms-staging", namespace="datahub-next"}
      values: '0'
    alert_rule_test:
    - alertname: DatahubNextServiceUnavailable
      eval_time: 5m
      exp_alerts:
      - exp_annotations:
          description: 'datahub-next is unavailable on k8s-dse'
          summary: 'datahub-next is unavailable on k8s-dse'
          dashboard: "https://grafana.wikimedia.org/d/hyl18XgMk/kubernetes-container-details?orgId=1&var-datasource=thanos&var-site=eqiad&var-cluster=k8s-dse&var-namespace=datahub-next&var-container=All"
          runbook: "https://wikitech.wikimedia.org/wiki/Data_Engineering/Systems/DataHub/Administration#Alerting"
        exp_labels:
          namespace: datahub-next
          deployment: datahub-gms-staging
          severity: warning
          team: data-platform
  - input_series:
    - series: kube_deployment_status_replicas_available{deployment="datahub-gms-production", namespace="datahub"}
      values: '0'
    alert_rule_test:
    - alertname: DatahubServiceUnavailable
      eval_time: 5m
      exp_alerts:
      - exp_annotations:
          description: 'datahub is unavailable on k8s-dse'
          summary: 'datahub is unavailable on k8s-dse'
          dashboard: "https://grafana.wikimedia.org/d/hyl18XgMk/kubernetes-container-details?orgId=1&var-datasource=thanos&var-site=eqiad&var-cluster=k8s-dse&var-namespace=datahub&var-container=All"
          runbook: "https://wikitech.wikimedia.org/wiki/Data_Engineering/Systems/DataHub/Administration#Alerting"
        exp_labels:
          namespace: datahub
          deployment: datahub-gms-production
          severity: critical
          team: data-platform
