rule_files:
  - druid.yaml
evaluation_interval: 1m
tests:

## Druid Realtime Ingestion tests
    - interval: 1m
      input_series:
        - series: druid_realtime_ingest_events_processed_count{cluster="druid_analytics", datasource="wmf_netflow", instance="an-druid1006:8000", job="druid", prometheus="analytics", site="eqiad"}
          values: '0'
        - series: druid_realtime_ingest_events_processed_count{cluster="druid_analytics", datasource="wmf_netflow", instance="an-druid1007:8000", job="druid", prometheus="analytics", site="eqiad"}
          values: '0'
        - series: druid_realtime_ingest_events_processed_count{cluster="druid_analytics", datasource="wmf_netflow", instance="an-druid1003:8000", job="druid", prometheus="analytics", site="eqiad"}
          values: '0'
        - series: druid_realtime_ingest_events_processed_count{cluster="druid_analytics", datasource="wmf_netflow", instance="an-druid1004:8000", job="druid", prometheus="analytics", site="eqiad"}
          values: '0'
        - series: druid_realtime_ingest_events_processed_count{cluster="druid_analytics", datasource="wmf_netflow", instance="an-druid1005:8000", job="druid", prometheus="analytics", site="eqiad"}
          values: '0'
      alert_rule_test:
      - alertname: DruidNetflowSupervisor
        eval_time: 5m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
          exp_annotations:
            description: There have been zero wmf_netflow events received on the druid_analytics cluster over a 30 minutes period.
            summary: Zero wmf_netflow events received by druid_analytics over the last 30 minutes.
            dashboard: https://grafana.wikimedia.org/d/000000538/druid?refresh=1m&var-cluster=druid_analytics&panelId=41&fullscreen&orgId=1
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Druid/Alerts#Druid_Netlow_Supervisor
        - exp_labels:
            severity: warning
            team: data-platform
          exp_annotations:
            description: There have been fewer than 10 wmf_netflow events received on the druid_analytics cluster over a 30 minutes period.
            summary: Fewer than 10 wmf_netflow events received by druid_analytics over the last 30 minutes.
            dashboard: https://grafana.wikimedia.org/d/000000538/druid?refresh=1m&var-cluster=druid_analytics&panelId=41&fullscreen&orgId=1
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Druid/Alerts#Druid_Netlow_Supervisor

## Druid Segments Unavailable tests
    - interval: 1m
      input_series:
        - series: druid_coordinator_segment_unavailable_count{cluster="druid_analytics", datasource="event_navigationtiming", instance="an-druid1006:8000", job="druid", prometheus="analytics", site="eqiad"}
          values: '11 11 11 11 11 11 11 11 11 11 11 11 11 11 11'
      alert_rule_test:
      - alertname: DruidSegmentsUnavailable
        eval_time: 5m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
            cluster: druid_analytics
            datasource: event_navigationtiming
          exp_annotations:
            description: More than 10 segments of the event_navigationtiming datastource have been reported as being unavailable on the druid_analytics Druid cluster, over a 15 minutes period.
            summary: More than 10 segments have been unavailable for event_navigationtiming on the druid_analytics Druid cluster.
            dashboard: https://grafana.wikimedia.org/d/000000538/druid?refresh=1m&var-cluster=druid_analytics&panelId=49&fullscreen&orgId=1
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Druid/Alerts#Druid_Segments_Unavailable
        - exp_labels:
            severity: warning
            team: data-platform
            cluster: druid_analytics
            datasource: event_navigationtiming
          exp_annotations:
            description: More than 5 segments of the event_navigationtiming datastource have been reported as being unavailable on the druid_analytics Druid cluster, over a 15 minutes period.
            summary: More than 5 segments have been unavailable for event_navigationtiming on the druid_analytics Druid cluster.
            dashboard: https://grafana.wikimedia.org/d/000000538/druid?refresh=1m&var-cluster=druid_analytics&panelId=49&fullscreen&orgId=1
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Druid/Alerts#Druid_Segments_Unavailable
