rule_files:
  - hadoop-yarn.yaml
evaluation_interval: 1m
tests:

## YARN Unhealthy Worker tests
    - interval: 1m
      input_series:
        - series: Hadoop_ResourceManager_NumUnhealthyNMs{hadoop_cluster="analytics-hadoop", instance="an-master1001:10083"}
          values: '5'
      alert_rule_test:
      - alertname: HadoopYarnUnhealthyWorkers
        eval_time: 5m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: 3 or more Yarn Nodemanagers on the analytics-hadoop cluster are in an unhealthy state.
            summary: 3 or more unhealthy Yarn Nodemanagers on analytics-hadoop.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=46&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#Unhealthy_Yarn_Nodemanagers
        - exp_labels:
            severity: warning
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: 1 or more Yarn Nodemanagers on the analytics-hadoop cluster are in an unhealthy state.
            summary: 1 or more unhealthy Yarn Nodemanagers on analytics-hadoop.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=46&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#Unhealthy_Yarn_Nodemanagers
