rule_files:
  - hadoop-hdfs.yaml
evaluation_interval: 1m
tests:

## HDFS Corrupt Blocks tests
    - interval: 1m
      input_series:
        - series: Hadoop_NameNode_CorruptBlocks{hadoop_cluster="analytics-hadoop", instance="an-master1001:10080"}
          values: '60'
      alert_rule_test:
      - alertname: HdfsCorruptBlocks
        eval_time: 5m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: More than 50 HDFS corrupt blocks detected on the analytics-hadoop HDFS cluster.
            summary: HDFS corrupt blocks detected on the analytics-hadoop HDFS cluster.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=39&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#HDFS_corrupt_blocks
        - exp_labels:
            severity: warning
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: More than 30 HDFS corrupt blocks detected on the analytics-hadoop HDFS cluster.
            summary: HDFS corrupt blocks detected on the analytics-hadoop HDFS cluster.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=39&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#HDFS_corrupt_blocks

## HDFS Missing Blocks tests
    - interval: 1m
      input_series:
        - series: Hadoop_NameNode_MissingBlocks{hadoop_cluster="analytics-hadoop", instance="an-master1001:10080"}
          values: '10'
      alert_rule_test:
      - alertname: HdfsMissingBlocks
        eval_time: 5m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: More than 5 HDFS missing blocks detected on the analytics-hadoop HDFS cluster.
            summary: HDFS missing blocks detected on the analytics-hadoop HDFS cluster.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=40&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#HDFS_missing_blocks
        - exp_labels:
            severity: warning
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: More than 2 HDFS missing blocks detected on the analytics-hadoop HDFS cluster.
            summary: HDFS missing blocks detected on the analytics-hadoop HDFS cluster.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=40&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#HDFS_missing_blocks

## HDFS Total Files tests
    - interval: 1m
      input_series:
        - series: Hadoop_NameNode_FilesTotal{hadoop_cluster="analytics-hadoop", instance="an-master1001:10080"}
          values: '110000001'
      alert_rule_test:
      - alertname: HdfsTotalFilesHeap
        eval_time: 5m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: The total number of files on the analytics-hadoop HDFS cluster are more than the heap can support.
            summary: Total files on the analytics-hadoop HDFS cluster are more than the heap can support.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=28&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#HDFS_total_files_and_heap_size
        - exp_labels:
            severity: warning
            team: data-platform
            hadoop_cluster: analytics-hadoop
          exp_annotations:
            description: The total number of files on the analytics-hadoop HDFS cluster are more than the heap can support.
            summary: Total files on the analytics-hadoop HDFS cluster are more than the heap can support.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&orgId=1&panelId=28&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/Cluster/Hadoop/Alerts#HDFS_total_files_and_heap_size

## HDFS Namenode FSImage Age
    - interval: 1m
      input_series:
        - series: Hadoop_NameNode_LastCheckpointTime{hadoop_cluster="analytics-hadoop", job="jmx_hadoop", name="FSNamesystem", instance="an-master1001:10080"}
          # timestamp for metrics start at 0 (https://github.com/prometheus/prometheus/blob/main/cmd/promtool/unittest.go#L179)
          # Keep our epoch checkpoint time at 0 for 120m (7200s) + 6m (5m duration for alert)
          values: '0x126'
      alert_rule_test:
      - alertname: HdfsFSImageAge
        eval_time: 126m
        exp_alerts:
        - exp_labels:
            severity: critical
            team: data-platform
            hadoop_cluster: analytics-hadoop
            instance: an-master1001:10080
            job: jmx_hadoop
            name: FSNamesystem
          exp_annotations:
            description: The HDFS FSImage on analytics-hadoop:an-master1001:10080 is too old.
            summary: The HDFS FSImage on analytics-hadoop:an-master1001:10080 is too old.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&viewPanel=129&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Cluster/Hadoop/Alerts#HDFS_FSImage_too_old
        - exp_labels:
            severity: warning
            team: data-platform
            hadoop_cluster: analytics-hadoop
            instance: an-master1001:10080
            job: jmx_hadoop
            name: FSNamesystem
          exp_annotations:
            description: The HDFS FSImage on analytics-hadoop:an-master1001:10080 is too old.
            summary: The HDFS FSImage on analytics-hadoop:an-master1001:10080 is too old.
            dashboard: https://grafana.wikimedia.org/d/000000585/hadoop?var-hadoop_cluster=analytics-hadoop&viewPanel=129&fullscreen
            runbook: https://wikitech.wikimedia.org/wiki/Analytics/Cluster/Hadoop/Alerts#HDFS_FSImage_too_old
