rule_files:
  - mediawiki-dumps-v1.yaml
evaluation_interval: 1m

tests:
  # Test 1: Normal space usage (above thresholds)
  - interval: 1m
    input_series:
      - series: 'kubelet_volume_stats_available_bytes{namespace="mediawiki-dumps-legacy", persistentvolumeclaim="mediawiki-dumps-legacy-fs"}'
        values: '100+0x12'  # Keep constant value
      - series: 'kubelet_volume_stats_capacity_bytes{namespace="mediawiki-dumps-legacy", persistentvolumeclaim="mediawiki-dumps-legacy-fs"}'
        values: '100+0x12'  # Keep capacity constant
    alert_rule_test:
      - eval_time: 5m
        alertname: MediawikiDumpsV1PVCAvailableSpaceWarning
        exp_alerts: []  # No alerts firing
      - eval_time: 5m
        alertname: MediawikiDumpsV1PVCAvailableSpaceCritical
        exp_alerts: []  # No alerts firing

  # Test 2: Space usage crossing warning threshold (86% full)
  - interval: 1m
    input_series:
      - series: 'kubelet_volume_stats_available_bytes{namespace="mediawiki-dumps-legacy", persistentvolumeclaim="mediawiki-dumps-legacy-fs"}'
        values: '14+0x12'  # Available space = 14 units
      - series: 'kubelet_volume_stats_capacity_bytes{namespace="mediawiki-dumps-legacy", persistentvolumeclaim="mediawiki-dumps-legacy-fs"}'
        values: '100+0x12'  # Capacity = 100 units (14% available = 86% full)
    alert_rule_test:
      - eval_time: 4m
        alertname: MediawikiDumpsV1PVCAvailableSpaceWarning
        exp_alerts: []
      - eval_time: 9m
        alertname: MediawikiDumpsV1PVCAvailableSpaceWarning
        exp_alerts:
          - exp_labels:
              severity: warning
              team: data-platform
              namespace: mediawiki-dumps-legacy
              persistentvolumeclaim: mediawiki-dumps-legacy-fs
            exp_annotations:
              summary: 'PVC mediawiki-dumps-legacy-fs.mediawiki-dumps-legacy is 85% full'
              description: 'PVC mediawiki-dumps-legacy-fs.mediawiki-dumps-legacy is 85% full'
              dashboard: "https://grafana.wikimedia.org/goto/Ge3TKCwNR?orgId=1"
              runbook: "https://wikitech.wikimedia.org/wiki/Dumps/Airflow/Operations#The_dumps_CephFS_volume_is_filling_up"

  # Test 3: Space usage crossing critical threshold (91% full)
  - interval: 1m
    input_series:
      - series: 'kubelet_volume_stats_available_bytes{namespace="mediawiki-dumps-legacy", persistentvolumeclaim="mediawiki-dumps-legacy-fs"}'
        values: '9+0x12'  # Available space = 9 units
      - series: 'kubelet_volume_stats_capacity_bytes{namespace="mediawiki-dumps-legacy", persistentvolumeclaim="mediawiki-dumps-legacy-fs"}'
        values: '100+0x12'  # Capacity = 100 units (9% available = 91% full)
    alert_rule_test:
      - eval_time: 4m
        alertname: MediawikiDumpsV1PVCAvailableSpaceCritical
        exp_alerts: []
      - eval_time: 9m
        alertname: MediawikiDumpsV1PVCAvailableSpaceCritical
        exp_alerts:
          - exp_labels:
              severity: critical
              team: data-platform
              namespace: mediawiki-dumps-legacy
              persistentvolumeclaim: mediawiki-dumps-legacy-fs
            exp_annotations:
              summary: 'PVC mediawiki-dumps-legacy-fs.mediawiki-dumps-legacy is 90% full'
              description: 'PVC mediawiki-dumps-legacy-fs.mediawiki-dumps-legacy is 90% full'
              dashboard: "https://grafana.wikimedia.org/goto/Ge3TKCwNR?orgId=1"
              runbook: "https://wikitech.wikimedia.org/wiki/Dumps/Airflow/Operations#The_dumps_CephFS_volume_is_filling_up"