rule_files:
  - benthos.yaml
evaluation_interval: 1m
tests:
 - interval: 1m
   input_series:
    - series: 'kafka_burrow_partition_lag{exported_cluster="logging-eqiad",group="benthos-eqiad"}'
      values: '5000+0x20'
   external_labels:
     site: eqiad
   alert_rule_test:
    - alertname: BenthosKafkaConsumerLag
      eval_time: 12m
      exp_alerts:
       - exp_labels:
           severity: critical
           team: o11y
           exported_cluster: logging-eqiad
           group: benthos-eqiad
         exp_annotations:
          description: Benthos consumers are not keeping up with the rate of log production.
          summary: Too many messages in logging-eqiad for group benthos-eqiad
          dashboard: https://grafana.wikimedia.org/d/000000484/kafka-consumer-lag?var-cluster=logging-eqiad&var-datasource=eqiad%20prometheus/ops&var-consumer_group=benthos-eqiad
          runbook: TODO
