rule_files:
  - cdn.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: trafficserver_backend_requests_seconds_count{backend="pinkunicorn.service",cluster="cache_upload",instance="cp2",job="trafficserver-upload",method="GET",prog="atsbackend.mtail",site="codfw",status="500"}
        values: '0+240x20'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: ATSBackendErrorsHigh
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              severity: page
              team: sre
              backend: 'pinkunicorn.service'
              cluster: 'cache_upload'
            exp_annotations:
              summary: 'ATS: elevated 5xx errors from pinkunicorn.service in codfw #page'
              description: 'High number of 5xx responses from pinkunicorn.service to ATS in codfw'
              runbook: https://wikitech.wikimedia.org/wiki/Apache_Traffic_Server#Debugging
              dashboard: https://grafana.wikimedia.org/d/1T_4O08Wk/ats-backends-origin-servers-overview?orgId=1&viewPanel=12&var-site=codfw&var-cluster=upload&var-origin=pinkunicorn.service
