rule_files:
  - etcd.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'etcdmirror_lag{cluster="etcd", instance="conf2005:8000", job="etcdmirror", prefix="/conftool", url="https://conf2005.codfw.wmnet:2379"}'
        values: '0 0 0 1+30x90'
    alert_rule_test:
      - alertname: EtcdReplicationLag
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              severity: page
              cluster: etcd
              instance: conf2005:8000
              prefix: /conftool
              team: sre
              url: https://conf2005.codfw.wmnet:2379
              job: etcdmirror
            exp_annotations:
              summary: 'etcd-mirror is lagging: 661 #page'
              description: 'The etcd-mirror process on conf2005:8000 has detected lag for https://conf2005.codfw.wmnet:2379'
              runbook: https://wikitech.wikimedia.org/wiki/Etcd
              dashboard: TODO

  - interval: 1m
    input_series:
      - series: 'up{job="etcdmirror", instance="conf2005:8000", cluster="etcd"}'
        values: '1 1 1 0 0'
    alert_rule_test:
      - alertname: EtcdReplicationDown
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              severity: page
              cluster: etcd
              instance: conf2005:8000
              team: sre
              job: etcdmirror
            exp_annotations:
              summary: 'etcd replication down on conf2005:8000 #page'
              description: 'The etcd-mirror instance on conf2005:8000 is down or unresponsive'
              runbook: https://wikitech.wikimedia.org/wiki/Etcd/Main_cluster#Replication
              dashboard: TODO

  - interval: 1m
    input_series:
      - series: 'etcd_http_received_total{job="etcd", cluster="etcd", instance="conf2005:2371", method="GET"}'
        values: '0+100000x60'
      - series: 'etcd_http_received_total{job="etcd", cluster="etcd", instance="conf2006:2371", method="PUT"}'
        values: '0+100000x60'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: EtcdTrafficHigh
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: page
              instance: conf2005:2371
              team: sre
            exp_annotations:
              summary: 'Excessive traffic on etcd conf2005:2371 #page'
              description: 'The etcd instance on conf2005:2371 is receiving excessive read traffic'
              runbook: https://wikitech.wikimedia.org/wiki/Etcd/Main_cluster
              dashboard: 'https://grafana.wikimedia.org/d/tTE9nvdMk/etcd?orgId=1&viewPanel=4&var-datasource=thanos&var-site=codfw&var-cluster=etcd'
