rule_files:
  - ganeti.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'node_memory_MemAvailable_bytes{cluster="ganeti",instance="ganeti1001:9100",job="node", site="eqiad"}'
        values: '3092376453x100'
      - series: 'node_memory_MemTotal_bytes{cluster="ganeti",instance="ganeti1001:9100",job="node", site="eqiad"}'
        values: '34359738368x100'
      - series: 'node_memory_MemAvailable_bytes{cluster="ganeti",instance="ganeti1002:9100",job="node", site="eqiad"}'
        values: '1073741824x100'
      - series: 'node_memory_MemTotal_bytes{cluster="ganeti",instance="ganeti1002:9100",job="node", site="eqiad"}'
        values: '34359738368x100'
      - series: 'remote_instance:gnmi_bgp_neighbor_session_state{instance="ganeti2033:0", peer_group="Ganeti4", remote_instance="lsw1-b7-codfw:9804", site="codfw"}'
        values: '1 1'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: GanetiMemoryPressure
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              team: infrastructure-foundations
              severity: warning
              job: node
              instance: ganeti1001:9100
              cluster: ganeti
              site: eqiad
            exp_annotations:
              description: 'Ganeti: High memory usage on ganeti1001:9100'
              summary: 'Ganeti: High memory usage (91%) on ganeti1001:9100'
              dashboard: 'https://grafana.wikimedia.org/d/gd6vep5Iz/ganeti-memory-pressure?orgId=1&var-site=eqiad'
              runbook: 'https://wikitech.wikimedia.org/wiki/Ganeti#Memory_pressure'
          - exp_labels:
              team: infrastructure-foundations
              severity: critical
              job: node
              instance: ganeti1002:9100
              cluster: ganeti
              site: eqiad
            exp_annotations:
              description: 'Ganeti: High memory usage on ganeti1002:9100'
              summary: 'Ganeti: High memory usage (96.88%) on ganeti1002:9100'
              dashboard: 'https://grafana.wikimedia.org/d/gd6vep5Iz/ganeti-memory-pressure?orgId=1&var-site=eqiad'
              runbook: 'https://wikitech.wikimedia.org/wiki/Ganeti#Memory_pressure'
      - alertname: GanetiBGPDown
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              peer_group: Ganeti4
              instance: ganeti2033:0
              remote_instance: lsw1-b7-codfw:9804
              team: sre
              severity: critical
              scope: network
              site: codfw
            exp_annotations:
              summary: "BGP session down between ganeti2033 and lsw1-b7-codfw - group Ganeti4"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/ed8da087-4bcb-407d-9596-d158b8145d45/bgp-neighbors-detail?orgId=1&var-site=codfw&var-device=lsw1-b7-codfw:9804&var-bgp_group=Ganeti4&var-bgp_neighbor=ganeti2033'
              runbook: 'https://wikitech.wikimedia.org/wiki/Ganeti#GanetiBGPDown'