rule_files:
  - hardware.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'device_smart_healthy{cluster="alerting",instance="alert1001:9100",device="sda"}'
        values: '0'
    alert_rule_test:
      - alertname: SmartNotHealthy
        eval_time: 0m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: sre
              instance: alert1001:9100
              cluster: alerting
              device: sda
            exp_annotations:
              summary: "Disk not healthy"
              description: >-
                The disk SMART status is *not* healthy, this could be an early
                warning before the disk fails.
              dashboard: 'https://grafana.wikimedia.org/d/000000377/host-overview?orgId=1&var-server=alert1001'
              runbook: https://wikitech.wikimedia.org/wiki/SMART#Alerts
  - interval: 1m
    input_series:
      - series: 'ipmi_temperature_state{instance="alert1001:9290",name="Temp"}'
        values: '1'
    alert_rule_test:
      - alertname: Temperature
        exp_alerts:
          - exp_labels:
              severity: critical
              team: sre
              instance: alert1001:9290
              name: "Temp"
            exp_annotations:
              summary: "Temp issue on alert1001:9290"
              description: "Temp issue on alert1001:9290"
              dashboard: "https://grafana.wikimedia.org/d/ZA1I-IB4z/ipmi-sensor-state?orgId=1&viewPanel=92&var-server=alert1001"
              runbook: https://wikitech.wikimedia.org/wiki/Dc-operations/Hardware_Troubleshooting_Runbook
  - interval: 1m
    input_series:
      - series: 'ipmi_sensor_state{instance="alert1001:9290",name="PS Redundancy",type="Power Supply"}'
        values: '1'
    alert_rule_test:
      - alertname: PowerSupply
        exp_alerts:
          - exp_labels:
              severity: critical
              team: sre
              instance: alert1001:9290
              name: "PS Redundancy"
              type: "Power Supply"
            exp_annotations:
              summary: "Power Supply - PS Redundancy - issue on alert1001:9290"
              description: "Power Supply - PS Redundancy - issue on alert1001:9290"
              dashboard: "https://grafana.wikimedia.org/d/ZA1I-IB4z/ipmi-sensor-state?orgId=1&var-Sensor=Power%20Supply&var-server=alert1001"
              runbook: https://wikitech.wikimedia.org/wiki/Dc-operations/Hardware_Troubleshooting_Runbook#Power_Supply_Failures
