rule_files:
  - ircstream.yaml
tests:
  - interval: 1m
    input_series:
      - series: 'ircstream_messages_total{channel="#aa.wikibooks", cluster="misc", instance="irc2001:16667", job="ircstream"}'
        values: '0+0x10'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: UdpIRCStreamThroughput
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: sre
              instance: "irc2001:16667"
            exp_annotations:
              summary: "irc2001:16667 has relayed less than 100 messages over past 5 minutes }}"
              description: "irc2001:16667 has relayed less than 100 messages for the irc.wikimedia.org service over past 5 minutes }}"
              dashboard: https://grafana.wikimedia.org/d/eb101795-c69e-4b9c-b848-f042d604f234/ircstream
              runbook: https://wikitech.wikimedia.org/wiki/Irc.wikimedia.org

