rule_files:
- jobqueue.yaml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  # Job1 and Job2 delays have a p95 over 120s.
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="1"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="30"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="60"}
    values: "0x30 0+54x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="300"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="+Inf"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job1"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="1"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="30"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="60"}
    values: "0x30 0+54x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="300"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="+Inf"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job2"}
    values: "0x30 0+60x380"
  # Job3 delays are always short.
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="1"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="30"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="+Inf"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job3"}
    values: "0x30 0+60x380"
  # Job4 delays have also been short historically, but it's also not currently active (is ignored).
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job4", le="1"}
    values: "0x410"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job4", le="30"}
    values: "60x410"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job4", le="+Inf"}
    values: "60x410"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job4"}
    values: "60x410"
  external_labels:
    site: antarctica
  alert_rule_test:
  - alertname: JobQueueLowTrafficConsumerWidespreadHighLatency
    eval_time: 360m
    exp_alerts: []  # Not firing (too short)
  - alertname: JobQueueLowTrafficConsumerWidespreadHighLatency
    eval_time: 400m
    exp_alerts:
    - exp_labels:
        severity: critical
        team: sre
      exp_annotations:
        summary: "Processing delay times for low-traffic consumer rules are unusually high"
        description: "Processing delay times for low-traffic consumer rules are unusually high in site antarctica for at least 50% of active job types for 6h."
        runbook: "https://wikitech.wikimedia.org/wiki/MediaWiki_JobQueue/Operations#JobQueueLowTrafficConsumerWidespreadHighLatency"
        dashboard: "https://grafana.wikimedia.org/d/fe130675-0c2d-4991-9dec-f54cf6a9c4d8/jobqueue-low-traffic-jobs?var-site=antarctica"
- interval: 1m
  input_series:
  # Job1 and Job2 delays are always short.
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="1"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="30"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job1", le="+Inf"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job1"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="1"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="30"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job2", le="+Inf"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job2"}
    values: "0x30 0+60x380"
  # Job3 delays have a p95 over 120s.
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="1"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="30"}
    values: "0x30 0x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="60"}
    values: "0x30 0+54x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="300"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_bucket{rule="low-traffic-jobs-mediawiki-job-Job3", le="+Inf"}
    values: "0x30 0+60x380"
  - series: cpjobqueue_normal_rule_processing_delay_count{rule="low-traffic-jobs-mediawiki-job-Job3"}
    values: "0x30 0+60x380"
  external_labels:
    site: antarctica
  alert_rule_test:
  - alertname: JobQueueLowTrafficConsumerWidespreadHighLatency
    eval_time: 400m
    exp_alerts: []  # Not firing (too few jobs with high processing delay)
