# https://prometheus.io/docs/prometheus/latest/configuration/unit_testing_rules/
# Test locally using `promtool test rules kubernetes_generic_test.yaml`
rule_files:
- kubernetes-generic.yaml
evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: 'up{instance="kubefake1001:9091", job="calico-felix"}'
    values: "1 0 0 0 0 0 0"
  - series: 'rsyslog_kubernetes_record_seen_total{instance="kubefake1001:9105"}'
    values: "1 0 0 0 0 0"
  - series: 'certmanager_certificate_ready_status{namespace="fakespace", name="fakename", condition="False"}'
    values: "0 1 1 1 1 1 1 1 1 1 1"
  - series: 'helm_release_status{namespace="pinkunicorn", name="main", status="pending-upgrade"}'
    values: "0 1 1 1 1 1 1 1 1 1 1 1"
  - series: 'helm_release_status{namespace="pinkunicorn", name="main", status="deployed"}'
    values: "1 1 1 1 1 1 1 1 1 1 1 0"
  - series: 'kubelet_runtime_operations_duration_seconds_bucket{job="k8s-node", instance="kubefakenode:1234", operation_type="fakeop", le="+Inf"}'
    values: "1+1x6"
  - series: 'kubelet_runtime_operations_duration_seconds_bucket{job="k8s-node", instance="kubefakenode:1234", operation_type="fakeop", le="0.7"}'
    values: "0+0x6"
  - series: 'apiserver_request_count{job="k8s-api", code="200"}'
    values: "100+1x6"
  - series: 'apiserver_request_count{job="k8s-api", code="0"}'
    values: "100+1x6"
  - series: 'apiserver_request_count{job="k8s-api", code="500"}'
    values: "1+5x6"
  - series: 'apiserver_request_duration_seconds_bucket{verb="GET", resource="shoes", le="+Inf"}'
    values: "1+1x6"
  - series: 'apiserver_request_duration_seconds_bucket{verb="GET", resource="shoes", le="+0.6"}'
    values: "0+0x6"
  - series: 'apiserver_request_duration_seconds_bucket{verb="GET", resource="secrets", le="+Inf"}'
    values: "1+1x6"
  - series: 'apiserver_request_duration_seconds_bucket{verb="GET", resource="secrets", le="+0.6"}'
    values: "0+0x6"
  - series: 'apiserver_request_duration_seconds_bucket{verb="LIST", resource="secrets", le="+Inf"}'
    values: "1+1x6"
  - series: 'apiserver_request_duration_seconds_bucket{verb="LIST", resource="secrets", le="+1.8"}'
    values: "0+0x6"
  external_labels:
    site: moon
    prometheus: eagle
  alert_rule_test:
  - alertname: KubernetesCalicoDown
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        severity: critical
        team: sre
        instance: "kubefake1001:9091"
        job: "calico-felix"
      exp_annotations:
        summary: "kubefake1001:9091 is not running calico-node Pod"
        description: "Prometheus has been unable to scrape metrics from calico on host kubefake1001:9091 job(calico-felix) for more than 5 minutes. Make sure the calico-node Pod is running on that node."
        runbook: https://wikitech.wikimedia.org/wiki/Calico#Operations
  - alertname: KubernetesRsyslogDown
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        severity: critical
        team: sre
        instance: "kubefake1001:9105"
      exp_annotations:
        summary: "rsyslog on kubefake1001:9105 is missing kubernetes logs"
        description: "Rsyslog mmkubernetes plugin on kubefake1001:9105 did not see any records in the last 5 minutes. This might mean that the plugin was terminated."
        runbook: https://wikitech.wikimedia.org/wiki/Kubernetes/Logging#Common_issues
  - alertname: CertManagerCertNotReady
    eval_time: 11m
    exp_alerts:
    - exp_labels:
        severity: critical
        team: sre
        namespace: fakespace
        name: fakename
        condition: "False"
      exp_annotations:
        summary: "Certificate fakespace/fakename is not in a ready state"
        description: "The certificate fakename in namespace fakespace is in a not ready state for at least 10min"
        runbook: https://wikitech.wikimedia.org/wiki/Kubernetes/cert-manager
  - alertname: HelmReleaseBadStatus
    eval_time: 12m
    exp_alerts:
    - exp_labels:
        team: sre
        severity: critical
        namespace: pinkunicorn
        name: main
        status: pending-upgrade
      exp_annotations:
        summary: "Helm release pinkunicorn/main on eagle@moon in state pending-upgrade"
        description: "The Helm release pinkunicorn/main on Kubernetes cluster eagle@moon is in state pending-upgrade for more than 11min, you might need to roll back"
        runbook: "https://wikitech.wikimedia.org/wiki/Kubernetes/Deployments#Rolling_back_in_an_emergency"
  - alertname: KubeletOperationalLatency
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        team: sre
        severity: warning
        job: k8s-node
        instance: kubefakenode:1234
        operation_type: fakeop
      exp_annotations:
        summary: "High latency for kubelet fakeop operation on kubefakenode:1234"
        description: "Kubelet fakeop operations on kubefakenode:1234 take 700ms in p99"
        runbook: "https://wikitech.wikimedia.org/wiki/Kubernetes"
        dashboard: "https://grafana.wikimedia.org/d/000000472/kubernetes-kubelets?var-cluster=moon prometheus/eagle"
  - alertname: KubernetesAPIErrorRate
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        team: sre
        severity: warning
        code: 500
      exp_annotations:
        description: "Kubernetes API servers eagle@moon are responding with 71.43% HTTP 500s"
        summary: "High rate of HTTP 500 responses from Kubernetes API eagle@moon"
        runbook: "https://wikitech.wikimedia.org/wiki/Kubernetes"
        dashboard: "https://grafana-rw.wikimedia.org/d/000000435/kubernetes-api?var-site=moon&var-cluster=eagle"
  - alertname: KubernetesAPILatency
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        team: sre
        severity: critical
        verb: GET
        resource: shoes
      exp_annotations:
        description: "Kubernetes API GET requests for shoes take 600ms in p99 on eagle@moon"
        summary: "High Kubernetes API latency (GET shoes) on eagle@moon"
        runbook: "https://wikitech.wikimedia.org/wiki/Kubernetes"
        dashboard: "https://grafana-rw.wikimedia.org/d/000000435/kubernetes-api?var-site=moon&var-cluster=eagle"
  - alertname: KubernetesAPILatencySecrets
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        team: sre
        severity: critical
        verb: GET
        resource: secrets
      exp_annotations:
        description: "Kubernetes API GET requests for secrets take 600ms in p99 on eagle@moon"
        summary: "High Kubernetes API latency (GET secrets) on eagle@moon"
        runbook: "https://wikitech.wikimedia.org/wiki/Kubernetes"
        dashboard: "https://grafana-rw.wikimedia.org/d/000000435/kubernetes-api?var-site=moon&var-cluster=eagle"
  - alertname: KubernetesAPILatencySecretsLIST
    eval_time: 6m
    exp_alerts:
    - exp_labels:
        team: sre
        severity: critical
        verb: LIST
        resource: secrets
      exp_annotations:
        description: "Kubernetes API LIST requests for secrets take 1.8s in p99 on eagle@moon"
        summary: "High Kubernetes API latency (LIST secrets) on eagle@moon"
        runbook: "https://wikitech.wikimedia.org/wiki/Kubernetes"
        dashboard: "https://grafana-rw.wikimedia.org/d/000000435/kubernetes-api?var-site=moon&var-cluster=eagle"
