rule_files:
  - liberica.yaml
evaluation_interval: 2m
tests:
  - interval: 2m
    external_labels:
      site: ulsfo
    input_series:
      - series: 'liberica_cp_match_fp{cluster="liberica", instance="lvs4008:9100"}'
        values: '0 0'
    alert_rule_test:
      - alertname: LibericaDiffFPCheck
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              cluster: liberica
              instance: lvs4008:9100
              severity: critical
              team: sre
            exp_annotations:
              summary: "Liberica instance lvs4008:9100 control plane status doesn't match with forwarding plane status"
              description: "At least one service shows discrepancies between control plane and forwarding plane on lvs4008:9100"
              dashboard: 'https://grafana.wikimedia.org/d/fa4de97a-7114-48c7-a91a-f56089ef554f/liberica?var-site=ulsfo&var-instance=lvs4008'
              runbook: 'https://wikitech.wikimedia.org/wiki/Liberica#LibericaDiffFPCheck'
  - interval: 1m
    external_labels:
      site: ulsfo
    input_series:
      # Should pass
      - series: 'node_file_age_timestamp_seconds_total{cluster="liberica", instance="lvs4010:9100", job="node", path="/etc/liberica/config.yaml", prometheus="ops", site="ulsfo"}'
        values: '1742226643 1742226643'
      - series: 'liberica_cp_configuration_reload_timestamp_seconds{cluster="liberica", instance="lvs4010:3003", job="liberica", prometheus="ops", result="ok", site="ulsfo"}'
        values: '1742226671 1742226671'
      # Should fail
      - series: 'node_file_age_timestamp_seconds_total{cluster="liberica", instance="lvs4009:9100", job="node", path="/etc/liberica/config.yaml", prometheus="ops", site="ulsfo"}'
        values: '1842219443 1842219443'
      - series: 'liberica_cp_configuration_reload_timestamp_seconds{cluster="liberica", instance="lvs4009:3003", job="liberica", prometheus="ops", result="ok", site="ulsfo"}'
        values: '1742226671 1742226671'
    alert_rule_test:
      - alertname: LibericaStaleConfig
        eval_time: 2m
        exp_alerts:
         - exp_labels:
             hostname: 'lvs4009'
             cluster: "liberica"
             severity: critical
             team: sre
           exp_annotations:
             summary: "Liberica instance lvs4009 is running a stale configuration"
             description: "Liberica control plane daemon isn't using the latest configuration deployed by puppet on lvs4009"
             dashboard: 'https://grafana.wikimedia.org/d/fa4de97a-7114-48c7-a91a-f56089ef554f/liberica?orgId=1&viewPanel=10&var-site=ulsfo&var-instance=lvs4009'
             runbook: 'https://wikitech.wikimedia.org/wiki/Liberica#LibericaStaleConfig'
  - interval: 1m
    external_labels:
      site: esams
    input_series:
      - series: 'liberica_etcd_errors_total{cluster="liberica", instance="lvs3008:3003", job="liberica", key="/conftool/v1/pools/esams/ncredir/nginx", prometheus="ops", site="esams"}'
        values: '0 10 20 30 40 50 60'
    alert_rule_test:
      - alertname: LibericaEtcdErrors
        eval_time: 6m
        exp_alerts:
          - exp_labels:
              instance: 'lvs3008:3003'
              key: '/conftool/v1/pools/esams/ncredir/nginx'
              severity: critical
              team: sre
            exp_annotations:
              summary: "Liberica instance lvs3008:3003 is experiencing etcd issues"
              description: "Liberica control plane daemon is failing to watch etcd key /conftool/v1/pools/esams/ncredir/nginx on lvs3008:3003"
              dashboard: 'https://grafana.wikimedia.org/d/fa4de97a-7114-48c7-a91a-f56089ef554f/liberica?orgId=1&var-site=esams&var-instance=lvs3008&viewPanel=11'
              runbook: 'https://wikitech.wikimedia.org/wiki/Liberica#LibericaEtcdErrors'
  - interval: 1m
    external_labels:
      site: eqsin
    input_series:
      - series: 'liberica_cp_unhealthy_pooled_realservers_total{instance="lvs5005:3003", service="upload-httpslb6_443", site="eqsin"}'
        values: '2 2 2 2 2'
    alert_rule_test:
      - alertname: LibericaUnhealthyRealserverPooled
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              instance: 'lvs5005:3003'
              service: 'upload-httpslb6_443'
              severity: critical
              team: sre
              site: eqsin
            exp_annotations:
              summary: "Liberica service upload-httpslb6_443 has 2 unhealthy realservers pooled on lvs5005:3003"
              description: "Liberica control plane is keeping one or more unhealthy realservers pooled, this is likely triggered by the depool threshold being enforced"
              dashboard: 'https://grafana.wikimedia.org/d/d70d14db-4a71-414d-8425-7a30d7127ca6/liberica-services?orgId=1&var-site=eqsin&var-service=upload-httpslb6_443&var-instance=lvs5005'
              runbook: 'https://wikitech.wikimedia.org/wiki/Liberica#LibericaUnhealthyRealserverPooled'
