rule_files:
  - mail.yaml
evaluation_interval: 1m
tests:
  - alert_rule_test:
    - alertname: InboundMXQueueHigh
      eval_time: 70m
      exp_alerts:
       - exp_labels:
           team: sre
           severity: page
           instance: mx-in1001:9154
         exp_annotations:
           description: 'The inbound postfix mail exchanger mx-in1001:9154 has a lot of messages in its queues. Check the runbook for further troubleshooting.'
           summary: 'MX host mx-in1001:9154 has many queued messages: 2380 #page'
           dashboard: 'https://grafana.wikimedia.org/d/h36Havfik/mail-postfix-servers'
           runbook: 'https://wikitech.wikimedia.org/wiki/Postfix'
    interval: 1m
    input_series:
      - series: 'postfix_showq_message_size_bytes_count{cluster="misc", instance="mx-in1001:9154", job="postfix", prometheus="ops", queue="active", site="eqiad"}'
        values: '490+10x90'
      - series: 'postfix_showq_message_size_bytes_count{cluster="misc", instance="mx-in1001:9154", job="postfix", prometheus="ops", queue="incoming", site="eqiad"}'
        values: '490+10x90'
  - alert_rule_test:
    - alertname: InboundMXQueueNoMetrics
      eval_time: 70m
      exp_alerts:
       - exp_labels:
           team: sre
           severity: critical
         exp_annotations:
           description: 'Metric, postfix_showq_message_size_bytes_count, used for the inbound postfix queue length cannot be found'
           summary: 'Queue length metrics not found'
           dashboard: 'https://grafana.wikimedia.org/d/h36Havfik/mail-postfix-servers'
           runbook: 'https://wikitech.wikimedia.org/wiki/Postfix'
    interval: 1m
  - alert_rule_test:
    - alertname: OutboundMXQueueHigh
      eval_time: 70m
      exp_alerts:
       - exp_labels:
           team: sre
           severity: page
           instance: mx-out1001:9154
         exp_annotations:
           description: 'The outbound postfix mail exchanger mx-out1001:9154 has a lot of messages in its queues. Check the runbook for further troubleshooting.'
           summary: 'MX host mx-out1001:9154 has many queued messages: 2380 #page'
           dashboard: 'https://grafana.wikimedia.org/d/h36Havfik/mail-postfix-servers'
           runbook: 'https://wikitech.wikimedia.org/wiki/Postfix'
    interval: 1m
    input_series:
      - series: 'postfix_showq_message_size_bytes_count{cluster="misc", instance="mx-out1001:9154", job="postfix", prometheus="ops", queue="active", site="eqiad"}'
        values: '490+10x90'
      - series: 'postfix_showq_message_size_bytes_count{cluster="misc", instance="mx-out1001:9154", job="postfix", prometheus="ops", queue="incoming", site="eqiad"}'
        values: '490+10x90'
  - alert_rule_test:
    - alertname: OutboundMXQueueNoMetrics
      eval_time: 70m
      exp_alerts:
       - exp_labels:
           team: sre
           severity: critical
         exp_annotations:
           description: 'Metric, postfix_showq_message_size_bytes_count, used for the outbound postfix queue length cannot be found'
           summary: 'Queue length metrics not found'
           dashboard: 'https://grafana.wikimedia.org/d/h36Havfik/mail-postfix-servers'
           runbook: 'https://wikitech.wikimedia.org/wiki/Postfix'
    interval: 1m
