rule_files:
  - mediawiki.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'phpfpm_statustext_processes{cluster="api_appserver", service="php7.4-fpm.service", state="active"}'
        values: '0+42x60'
      - series: 'phpfpm_statustext_processes{cluster="api_appserver", service="php7.4-fpm.service", state="idle"}'
        values: '0+1x60'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: PHPFPMTooBusy
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              team: sre
              severity: page
              cluster: api_appserver
              service: php7.4-fpm.service
            exp_annotations:
              description: 'The MediaWiki cluster api_appserver in eqiad is experiencing saturation of php7.4-fpm.service workers 2.326%'
              summary: 'Not enough idle php7.4-fpm.service workers for Mediawiki api_appserver at eqiad #page'
              dashboard: 'https://grafana.wikimedia.org/d/RIA1lzDZk/application-servers-red-dashboard?from=now-3h&orgId=1&to=now&var-cluster=api_appserver&var-site=eqiad&viewPanel=64'
              runbook: 'https://bit.ly/wmf-fpmsat'

  - interval: 1m
    input_series:
      - series: 'phpfpm_statustext_processes{cluster="jobrunner", state="active", service="php7.2-fpm.service"}'
        values: '0+100x60'
      - series: 'phpfpm_statustext_processes{cluster="jobrunner", state="idle", service="php7.2-fpm.service"}'
        values: '0+1x60'
      - series: 'phpfpm_statustext_processes{cluster="jobrunner", state="active", service="php7.4-fpm.service"}'
        values: '0+100x60'
      - series: 'phpfpm_statustext_processes{cluster="jobrunner", state="idle", service="php7.4-fpm.service"}'
        values: '0+60x60'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: JobrunnerPHPBusyWorkers
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              team: sre
              severity: critical
              cluster: jobrunner
              service: php7.2-fpm.service
            exp_annotations:
              dashboard: 'https://grafana.wikimedia.org/d/wqj6s-unk/jobrunners?fullscreen&orgId=1&from=now-3h&to=now&var-datasource=eqiad%20prometheus/ops&viewPanel=54'
              description: 'Jobrunners in eqiad is experiencing saturation of php-fpm workers: 0.9901% for more than 10 minutes.'
              runbook: 'https://bit.ly/wmf-fpmsat'
              summary: 'Not enough idle php7.2-fpm.service workers for Mediawiki jobrunner at eqiad'

  - interval: 1m
    input_series:
      - series: 'up{cluster="api_appserver", instance="mw1317:9117", job="apache"}'
        values: '0+0x60'
      - series: 'up{cluster="api_appserver", instance="mw1318:9117", job="apache"}'
        values: '0+0x60'
      - series: 'up{cluster="api_appserver", instance="mw1319:9117", job="apache"}'
        values: '1+0x60'
      - series: 'up{cluster="api_appserver", instance="mw1320:9117", job="apache"}'
        values: '1+0x60'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: AppserversUnreachable
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              team: sre
              severity: critical
              cluster: api_appserver
            exp_annotations:
              description: 'Prometheus in eqiad is unable to scrape metrics for 50% of cluster api_appserver.'
              summary: 'Appserver unavailable for cluster api_appserver at eqiad'
              dashboard: 'https://grafana.wikimedia.org/d/RIA1lzDZk/application-servers-red-dashboard?orgId=1&var-site=eqiad&var-cluster=api_appserver'
              runbook: 'https://wikitech.wikimedia.org/wiki/Application_servers'

  # Test case for appservers GET
  - interval: 1m
    input_series:
      - series: 'cluster_code_method_handler:mediawiki_http_requests_duration:avg2m{cluster="appserver",method="GET", handler="proxy:unix:test", code="200"}'
        values: '1+0x10'
      - series: 'mediawiki_http_requests_duration_count{cluster="appserver",code="200",handler="proxy:unix:test",instance="mw1:3903",job="mtail",method="GET",prog="apache2-mediawiki.mtail",site="eqiad"}'
        values: '0+10000x10'
    external_labels:
      site: eqiad
    alert_rule_test:
      - alertname: MediaWikiLatencyExceeded
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              team: sre
              code: '200'
              handler: 'proxy:unix:test'
              method: 'GET'
              severity: critical
              cluster: appserver
            exp_annotations:
              description: 'The mediawiki cluster appserver in eqiad is experiencing slowness. The average latency for GET/200: 1s'
              summary: 'Average latency high: eqiad appserver GET/200'
              dashboard: 'https://grafana.wikimedia.org/d/RIA1lzDZk/application-servers-red-dashboard?panelId=9&fullscreen&orgId=1&from=now-3h&to=now&var-site=eqiad&var-cluster=appserver&var-method=GET'
              runbook: 'https://wikitech.wikimedia.org/wiki/Application_servers/Runbook#Average_latency_exceeded'

  # Test case for memcached errors
  - interval: 1m
    input_series:
      - series: 'log_mediawiki_servergroup_level_channel_doc_count{channel="memcached",cluster="logstash",instance="logstash2:9206",job="es_exporter",level="ERROR",servergroup="parsoid",site="codfw"}'
        values: '2000+1x10'
      - series: 'log_mediawiki_servergroup_level_channel_doc_count{channel="memcached",cluster="logstash",instance="logstash2:9206",job="es_exporter",level="ERROR",servergroup="appserver",site="codfw"}'
        values: '3400+1x10'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: MediaWikiMemcachedHighErrorRate
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              team: sre
              severity: critical
            exp_annotations:
              description: 'Elevated global rate of errors connecting to memcached from MediaWiki: 5420 errors/min'
              summary: 'MediaWiki memcached error rate is elevated globally'
              dashboard: 'https://grafana.wikimedia.org/d/000000438/mediawiki-exceptions-alerts?var-datasource=codfw%20prometheus/ops&viewPanel=19'
              runbook: 'https://wikitech.wikimedia.org/wiki/Application_servers/Runbook'

  # Test case for mediawiki fatals/errors
  - interval: 1m
    input_series:
      - series: 'log_mediawiki_servergroup_level_channel_doc_count{channel="fatal",cluster="logstash",instance="logstash2:9206",job="es_exporter",level="ERROR",servergroup="appserver",site="codfw"}'
        values: '20+1x10'
      - series: 'log_mediawiki_servergroup_level_channel_doc_count{channel="exception",cluster="logstash",instance="logstash2:9206",job="es_exporter",level="ERROR",servergroup="appserver",site="codfw"}'
        values: '85+1x10'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: MediaWikiHighErrorRate
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              team: sre
              severity: critical
              servergroup: appserver
            exp_annotations:
              description: 'Elevated global rate of errors and fatals from MediaWiki in server group appserver: 125/min'
              summary: 'Elevated rate of MediaWiki errors - appserver'
              dashboard: 'https://grafana.wikimedia.org/d/000000438/mediawiki-exceptions-alerts?panelId=18&fullscreen&orgId=1&var-datasource=codfw%20prometheus/ops'
              runbook: 'https://wikitech.wikimedia.org/wiki/Application_servers/Runbook'
