# XXX make alert fire in tests
#rule_files:
#  - node.yaml
#evaluation_interval: 1m
#tests:
#  - interval: 1m
#    input_series:
#      - series: 'node_textfile_mtime_seconds{file="metrics.prom",instance="instance1:9100"}'
#        values: '0+60x5 0+0x64'
#    alert_rule_test:
#      - alertname: NodeTextfileStale
#        eval_time: 5m
#      - alertname: NodeTextfileStale
#        eval_time: 70m
#        exp_alerts:
#          - exp_labels:
#              severity: critical
#              team: sre
#              instance: "instance1:9100"
#            exp_annotations:
#              summary: "Stale textfile for instance1:9100"
#              description: "Stale textfile for a"
#              dashboard: https://grafana.wikimedia.org/d/knkl4dCWz/node-exporter-textfile
#              runbook: https://wikitech.wikimedia.org/wiki/Prometheus#Stale_file_for_node-exporter_textfile
