rule_files:
  - oom.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'node_vmstat_oom_kill{cluster="api_appserver",instance="mw2001:9100",job="node",site="codfw"}'
        values: '0+1x100'
      - series: 'node_vmstat_oom_kill{cluster="kubernetes",instance="kubernetes2001:9100",job="node",site="codfw"}'
        values: '0+1x100'
      - series: 'node_vmstat_oom_kill{cluster="api_appserver",instance="mw2002:9100",job="node",site="codfw"}'
        values: '0x100'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: OOM
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              team: sre
              severity: warning
              job: node
              instance: mw2001:9100
              cluster: api_appserver
              site: codfw
            exp_annotations:
              description: 'OOM killer active on mw2001:9100'
              summary: 'OOM killer active on mw2001:9100'
              dashboard: 'https://grafana.wikimedia.org/d/-OcleDKIz/oom-kill'
              runbook: 'TODO'
