rule_files:
  - pki.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'x509_cert_expiry{cert="/etc/ssl/signers/aux/ca/aux.pem", cluster="pki", instance="pki2002:9100", job="node", site="codfw", subject="aux"}'
        values: '2419200-60x43200'
      - series: 'x509_cert_expiry{cert="/etc/ssl/signers/network/ca/network.pem", cluster="pki", instance="pki2002:9100", job="node", site="codfw", subject="network"}'
        values: '605000-60x43200'
    external_labels:
      site: codfw
    alert_rule_test:
      - alertname: PKICertificateExpiry
        eval_time: 4h
        exp_alerts:
          - exp_labels:
              team: infrastructure-foundations
              cert: /etc/ssl/signers/aux/ca/aux.pem
              job: node
              site: codfw
              severity: warning
              subject: aux
              cluster: pki
              instance: pki2002:9100
            exp_annotations:
              description: 'Intermediate certificate in the trust chain for aux will expire in less than 30 days'
              summary: 'Intermediate certificate in the trust chain for aux expires in 27d 16h 0m 0s'
              dashboard: 'TODO'
              runbook: 'https://wikitech.wikimedia.org/wiki/PKI/CA_Operations'
          - exp_labels:
              team: infrastructure-foundations
              cert: /etc/ssl/signers/network/ca/network.pem
              job: node
              site: codfw
              severity: critical
              subject: network
              cluster: pki
              instance: pki2002:9100
            exp_annotations:
              description: 'Intermediate certificate in the trust chain for network will expire in less than 7 days'
              summary: 'Intermediate certificate in the trust chain for network expires in 6d 16h 3m 20s'
              dashboard: 'TODO'
              runbook: 'https://wikitech.wikimedia.org/wiki/PKI/CA_Operations'
