rule_files:
  - proxoid.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: nginx_http_requests_total{site="eqiad", instance="urldownloader1001", status_code="500", cluster="hcaptcha"}
        values: '0+10x600'
      - series: nginx_http_requests_total{site="eqiad", instance="urldownloader1001", status_code="200", cluster="hcaptcha"}
        values: '0+900x600'
    external_labels:
      site: eqiad
    alert_rule_test:
      - eval_time: 4m
        alertname: ProxoidHigh5xxErrorRate
        exp_alerts: []   # no alert yet (still pending)
      - eval_time: 6m
        alertname: ProxoidHigh5xxErrorRate
        exp_alerts:
          - exp_labels:
              instance: urldownloader1001
              severity: critical
              team: sre
            exp_annotations:
              description: 'High 5xx error rate (>1%) on urldownloader1001 for eqiad'
              summary: 'urldownloader1001: >1% of requests are 5xx errors in eqiad'
              dashboard: 'https://grafana.wikimedia.org/d/441b2def-52e9-49d6-acad-91f5bb748989/hcaptcha-reverse-proxy?orgId=1&var-site=eqiad&var-instance=urldownloader1001'
              runbook: 'https://wikitech.wikimedia.org/wiki/HCaptcha#Runbook'
