rule_files:
  - puppet-agent.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'puppet_agent_failed{cluster="misc"}'
        values: '10+0x666'
      - series: 'puppet_agent_failed{cluster="misc"}'
        values: '100+0x666'
      - series: 'cluster:puppet_agent_resources_total:count0{cluster="misc"}'
        values: '10+0x666'
      - series: 'cluster:puppet_agent_resources_total:count{cluster="misc"}'
        values: '100+0x666'
    external_labels:
      site: eqiad
    alert_rule_test:
    - alertname: WidespreadPuppetFailure
      eval_time: 1m
      exp_alerts:
       - exp_labels:
           severity: critical
           team: sre
         exp_annotations:
           description: 'Widespread puppet agent failure in eqiad'
           summary: 'Puppet has failed in eqiad'
           dashboard: 'https://grafana.wikimedia.org/d/yOxVDGvWk/puppet?orgId=1&viewPanel=6'
           runbook: https://puppetboard.wikimedia.org/nodes?status=failed

  - interval: 1m
    input_series:
      - series: 'cluster:puppet_agent_failed:sum{cluster="misc"}'
        values: '10+0x666'
      - series: 'cluster:puppet_agent_failed:count{cluster="misc"}'
        values: '100+0x666'
      - series: 'cluster:puppet_agent_resources_total:count0{cluster="misc"}'
        values: '10+0x666'
      - series: 'cluster:puppet_agent_resources_total:count{cluster="misc"}'
        values: '100+0x666'
    alert_rule_test:
    - alertname: WidespreadPuppetFailure
      eval_time: 100m
      exp_alerts:
       - exp_labels:
           severity: warning
           team: sre
           cluster: misc
         exp_annotations:
           description: 'Widespread puppet agent failure on misc cluster'
           summary: 'Puppet has failed on misc cluster'
           dashboard: 'https://grafana.wikimedia.org/d/yOxVDGvWk/puppet?orgId=1&viewPanel=3&var-cluster=misc'
           runbook: https://puppetboard.wikimedia.org/nodes?status=failed

  - interval: 1m
    input_series:
      - series: 'role_owner{instance="alert1001:9100",team="infrastructure-foundation"}'
        values: '1x1441'
      - series: 'puppet_agent_failed{cluster="alerting",instance="alert1001:9100"}'
        values: '0+1x90'
    alert_rule_test:
    - alertname: PuppetFailure
      eval_time: 90m
      exp_alerts:
       - exp_labels:
           severity: critical
           team: infrastructure-foundation
           instance: alert1001:9100
           cluster: alerting
         exp_annotations:
           summary: Puppet has failed on alert1001:9100
           description: 'Puppet failure on alert1001:9100'
           dashboard: https://grafana.wikimedia.org/d/yOxVDGvWk/puppet
           runbook: https://puppetboard.wikimedia.org/nodes?status=failed

  - interval: 1m
    input_series:
      - series: 'role_owner{instance="alert1001:9100",team="infrastructure-foundation"}'
        values: '1x1441'
      - series: 'puppet_agent_resources_total{cluster="alerting",instance="alert1001:9100"}'
        values: '0x90'
    alert_rule_test:
    - alertname: PuppetZeroResources
      eval_time: 90m
      exp_alerts:
       - exp_labels:
           severity: warning
           team: infrastructure-foundation
           instance: alert1001:9100
           cluster: alerting
         exp_annotations:
           summary: Puppet has failed generate resources on alert1001:9100
           description: 'Zero Puppet resources on alert1001:9100'
           dashboard: https://grafana.wikimedia.org/d/yOxVDGvWk/puppet
           runbook: https://puppetboard.wikimedia.org/nodes?status=failed

  - interval: 1m
    input_series:
      - series: 'role_owner{instance="alert1001:9100",team="observability"}'
        values: '1x1441'
      - series: 'puppet_agent_resources_changed{cluster="alerting",instance="alert1001:9100",team="observability"}'
        values: '1x700 0 1x800'  # Ensure the alert doesn't trigger if we have one report with 0 changes
    alert_rule_test:
    - alertname: PuppetConstantChange
      eval_time: 1d
      exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'role_owner{instance="alert1001:9100",team="observability"}'
        values: '1x1441'
      - series: 'puppet_agent_enabled{cluster="alerting",instance="alert1001:9100"}'
        values: '1x10081'
      - series: 'puppet_agent_resources_changed{cluster="alerting",instance="alert1001:9100",team="observability"}'
        values: '1x1441'
    alert_rule_test:
    - alertname: PuppetConstantChange
      eval_time: 1d
      exp_alerts:
       - exp_labels:
           severity: warning
           team: observability
           instance: alert1001:9100
           cluster: alerting
         exp_annotations:
           summary: Puppet performing a change on every puppet run on alert1001:9100
           description: 'Puppet performing a change on every puppet run alert1001:9100'
           dashboard: https://grafana.wikimedia.org/d/yOxVDGvWk/puppet
           runbook: https://puppetboard.wikimedia.org/nodes?status=changed
  - interval: 1m
    input_series:
      - series: 'role_owner{instance="alert1001:9100",team="observability"}'
        values: '1x1441'
      - series: 'puppet_agent_resources_changed{cluster="alerting",instance="alert1001:9100",team="observability"}'
        values: '1x1441'
      - series: 'puppet_agent_enabled{cluster="alerting",instance="alert1001:9100"}'
        values: '0x10081'
    alert_rule_test:
    - alertname: PuppetConstantChange
      eval_time: 1d
      exp_alerts: []

  - interval: 1m
    input_series:
      - series: 'role_owner{instance="alert1001:9100",team="observability"}'
        values: '1x10081'
      - series: 'puppet_agent_enabled{cluster="alerting",instance="alert1001:9100"}'
        values: '0x10081'
    alert_rule_test:
    - alertname: PuppetDisabled
      eval_time: 1w
      exp_alerts:
       - exp_labels:
           severity: critical
           instance: alert1001:9100
           cluster: alerting
           team: observability
         exp_annotations:
           summary: Puppet disabled on alert1001:9100
           description: Puppet disabled on alert1001:9100
           dashboard: https://grafana.wikimedia.org/d/yOxVDGvWk/puppet?var-cluster=alerting&viewPanel=14
           runbook: https://wikitech.wikimedia.org/wiki/Puppet/Runbooks#Puppet_Disabled
