rule_files:
  - pybal.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'pybal_monitor_status{cluster="lvs", host="aqs1010.eqiad.wmnet", instance="lvs1019:9090", job="pybal", monitor="IdleConnection", service="aqs_7232"}'
        values: '1 0+0x1000'
      - series: 'pybal_monitor_status{cluster="lvs", host="cloudelastic1002.wikimedia.org", instance="lvs1020:9090", job="pybal", monitor="IdleConnection", service="cloudelasticlb_8243"}'
        values: '1 1+0x1000'
      - series: 'pybal_monitor_down_results_total{cluster="lvs",host="aqs1010.eqiad.wmnet",instance="lvs1019:9090",job="pybal",monitor="IdleConnection",service="aqs_7232"}'
        values: '1 1+1x1000'
    alert_rule_test:
      - alertname: PybalBackendDown
        eval_time: 13h
        exp_alerts:
          - exp_labels:
              severity: warning
              team: sre
              instance: aqs1010:0
              service: aqs_7232
            exp_annotations:
              summary: 'Pybal backend aqs1010:0 is down (aqs_7232)'
              description: 'Pybal has been failing health checks for aqs1010:0 for a long time.'
              runbook: TODO
              dashboard: TODO
  - interval: 1m
    external_labels:
      site: eqiad
    input_series:
      # Should trigger a fail
      - series: 'pybal_bgp_session_established{cluster="lvs", instance="lvs1234:9090", job="pybal", local_asn="64600", peer="192.0.2.1"}'
        values: '1 0 0 0 0 0 0'
      - series: 'pybal_bgp_enabled{cluster="lvs", instance="lvs1234:9090", job="pybal"}'
        values: '1 1 1 1 1 1 1'
      # Should pass
      - series: 'pybal_bgp_session_established{cluster="lvs", instance="lvs5678:9090", job="pybal", local_asn="64600", peer="192.0.2.2"}'
        values: '1 1 1 1 1 1 1'
      - series: 'pybal_bgp_enabled{cluster="lvs", instance="lvs5678:9090", job="pybal"}'
        values: '1 1 1 1 1 1 1'
    alert_rule_test:
      - alertname: PyBalBGPUnstable
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              cluster: "lvs"
              job: "pybal"
              instance: "lvs1234:9090"
              local_asn: "64600"
              peer: "192.0.2.1"
              severity: page
              team: sre
            exp_annotations:
              summary: "PyBal BGP sessions on instance lvs1234 with peer 192.0.2.1 are failing #page"
              description: "PyBal BGP session establishment with the peer 192.0.2.1 is unstable on host lvs1234."
              dashboard: "https://grafana.wikimedia.org/d/000000488/pybal-bgp?var-datasource=eqiad%20prometheus/ops&var-server=lvs1234"
              runbook: 'https://wikitech.wikimedia.org/wiki/PyBal#Alerts'
