rule_files:
  - systemd.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'systemd_unit_owner{instance="sretest1001:9100",team="serviceops",name="override.service"}'
        values: '1'
      - series: 'role_owner{instance="sretest1001:9100",team="infrastructure-foundations"}'
        values: '1'
      - series: 'instance_name:node_systemd_unit_state_failed:count1{instance="sretest1001:9100",name="example.service"}'
        values: '1'
      - series: 'instance_name:node_systemd_unit_state_failed:count1{instance="sretest1001:9100",name="override.service"}'
        values: '1'
    alert_rule_test:
      - alertname: SystemdUnitFailed
        exp_alerts:
        - exp_labels:
            severity: critical
            team: infrastructure-foundations
            instance: sretest1001:9100
            name: example.service
          exp_annotations:
            summary: "example.service Failed on sretest1001:9100"
            description: "example.service Failed on sretest1001:9100"
            dashboard: 'https://grafana.wikimedia.org/d/g-AaZRFWk/systemd-status'
            runbook: https://wikitech.wikimedia.org/wiki/Monitoring/check_systemd_state
        - exp_labels:
            severity: critical
            team: serviceops
            instance: sretest1001:9100
            name: override.service
          exp_annotations:
            summary: "override.service Failed on sretest1001:9100"
            description: "override.service Failed on sretest1001:9100"
            dashboard: 'https://grafana.wikimedia.org/d/g-AaZRFWk/systemd-status'
            runbook: https://wikitech.wikimedia.org/wiki/Monitoring/check_systemd_state

  - interval: 1m
    input_series:
      - series: 'systemd_unit_owner{instance="sretest1001:9100",team="serviceops",name="override.service"}'
        values: '1+0x60'
      - series: 'role_owner{instance="sretest1001:9100",team="infrastructure-foundations"}'
        values: '1+0x60'
      - series: 'node_systemd_service_restart_total{instance="sretest1001:9100",name="example.service"}'
        values: '5+10x60'
      - series: 'node_systemd_service_restart_total{instance="sretest1001:9100",name="override.service"}'
        values: '5+10x60'
    alert_rule_test:
      - alertname: SystemdUnitCrashLoop
        eval_time: 13m
        exp_alerts:
        - exp_labels:
            severity: warning
            team: infrastructure-foundations
            instance: sretest1001:9100
            name: example.service
          exp_annotations:
            summary: "example.service crashloop on sretest1001:9100"
            description: "example.service has had 100 restarts in the last 10m"
            dashboard: 'https://grafana.wikimedia.org/d/g-AaZRFWk/systemd-status'
            runbook: TODO
        - exp_labels:
            severity: warning
            team: serviceops
            instance: sretest1001:9100
            name: override.service
          exp_annotations:
            summary: "override.service crashloop on sretest1001:9100"
            description: "override.service has had 100 restarts in the last 10m"
            dashboard: 'https://grafana.wikimedia.org/d/g-AaZRFWk/systemd-status'
            runbook: TODO
