rule_files:
  - hardware.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'ipmi_sensor_state{instance="alert1001:9290",name="PS Redundancy",type="Power Supply"}'
        values: '1+0x600'
    alert_rule_test:
      - alertname: PowerSupplyFailure
        eval_time: 3h
        exp_alerts:
          - exp_labels:
              severity: task
              team: dcops
              instance: alert1001:9290
              name: "PS Redundancy"
              type: "Power Supply"
            exp_annotations:
              summary: "Power Supply - PS Redundancy - issue on alert1001:9290"
              description: "Power Supply - PS Redundancy - issue on alert1001:9290"
              dashboard: "https://grafana.wikimedia.org/d/ZA1I-IB4z/ipmi-sensor-state?orgId=1&var-Sensor=Power%20Supply&var-server=alert1001"
              runbook: https://wikitech.wikimedia.org/wiki/Dc-operations/Hardware_Troubleshooting_Runbook#Power_Supply_Failures

