rule_files:
  - eventlogging.yaml
tests:
  - interval: 1m
    input_series:
      - series: kafka_server_BrokerTopicMetrics_MessagesIn_total{instance="kafka-jumbo1001:7800",cluster="kafka_jumbo",topic="eventlogging_EventError"}
        values: "0+2000x30"
    alert_rule_test:
      - alertname: EventloggingErrorsHigh
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: data-engineering
            exp_annotations:
              description: Median event logging throughput above 30 events/s over past 15 minutes
              summary: Event logging throughput above 30 events/s over past 15 minutes
              dashboard: https://grafana.wikimedia.org/d/000000505/eventlogging?orgId=1&viewPanel=13
              runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/EventLogging/Administration

  - interval: 1m
    input_series:
      - series: kafka_server_BrokerTopicMetrics_MessagesIn_total{cluster="kafka_jumbo",instance="kafka-jumbo1007:7800",topic="eventlogging_NavigationTiming"}
        values: "0+0x30"
    alert_rule_test:
      - alertname: EventloggingNavigationTimingEventsLow
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: data-engineering
            exp_annotations:
              description: p50 of Navigation Timing event throughput has been under 1 req/sec over past 15m
              summary: p50 of Navigation Timing event throughput under 1 req/sec over past 15m
              dashboard: https://grafana.wikimedia.org/d/000000505/eventlogging?orgId=1&viewPanel=6
              runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/EventLogging/Administration

  - interval: 1m
    input_series:
      - series: kafka_server_BrokerTopicMetrics_MessagesIn_total{cluster="kafka_jumbo",instance="kafka-jumbo1001:7800",topic="eventlogging-client-side"}
        values: "0+500000x30"
    alert_rule_test:
      - alertname: EventloggingEventsHigh
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: data-engineering
            exp_annotations:
              description: Event logging p50 of overall event throughput above 1500 events/s for past 15 min
              summary: p50 of overall event throughput above 1500 events/s
              dashboard: https://grafana.wikimedia.org/d/000000505/eventlogging?orgId=1&viewPanel=6
              runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/EventLogging/Administration

  - interval: 1m
    input_series:
      - series: kafka_burrow_partition_lag{exported_cluster="jumbo-eqiad",group="eventlogging_processor_client_side_00",instance="kafkamon1002:9700",topic="eventlogging-client-side"}
        values: "0+150000x30"
    alert_rule_test:
      - alertname: EventLoggingKafkaLag
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: data-engineering
            exp_annotations:
              description: Kafka consumer lag for event logging over threshold for past 15 min, event logging may not be processing events, or has slowed down.
              summary: Kafka consumer lag for event logging over threshold for past 15 min.
              dashboard: https://grafana.wikimedia.org/d/000000484/kafka-consumer-lag?orgId=1&prometheus=ops&var-cluster=jumbo-eqiad&var-topic=All&var-consumer_group=eventlogging_processor_client_side_00
              runbook: https://wikitech.wikimedia.org/wiki/Analytics/Systems/EventLogging/Administration
