rule_files:
  - kafka-replication-factor.yaml
evaluation_interval: 1m
tests:
  - input_series:
    - series: kafka_cluster_Partition_ReplicasCount{topic="webrequest_text", kafka_cluster="jumbo-eqiad"}
      values: '1'
    alert_rule_test:
    - alertname: KafkaReplicationFactorTooLow
      exp_alerts:
      - exp_labels:
          severity: warning
          team: data-engineering
          topic: webrequest_text
          kafka_cluster: jumbo-eqiad
        exp_annotations:
          description: Kafka topic replication factor is too low
          summary: SKafka topic replication factor is too low
          dashboard: "https://grafana.wikimedia.org/d/000000234/kafka-by-topic?var-kafka_cluster=jumbo-eqiad&var-kafka_broker=All&var-topic=webrequest_text&viewPanel=40"
          runbook: "https://wikitech.wikimedia.org/wiki/Kafka/Administration#Increase_a_topic's_replication_factor"
