rule_files:
  - varnishkafka.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      # This should fire a warning because varnishkafka is relaying incoming requests at a rate of 15% of their incoming rate (20 ps / 3 ps) for at least 5 min
      - series: varnish_requests{instance="cp3050:3903", cluster="cache_text"}
        values: '0+20x11'
      - series: rdkafka_producer_topic_partition_msgs_total{instance="cp3050:9132", cluster="cache_text"}
        values: '0+3x8 200+20x2'
      # This should fire both a warning and a critical alert because varnishkafka is relaying incoming requests at a rate of only 5% of their incoming rate (20 ps / 1 ps) for at least 5 min
      - series: varnish_requests{instance="cp3051:3903", cluster="cache_text"}
        values: '0+20x11'
      - series: rdkafka_producer_topic_partition_msgs_total{instance="cp3051:9132", cluster="cache_text"}
        values: '0+1x10 240'
      # This should not fire an alert because it simulates a depooled server, where varnish requests and varnishkafka messages both drop to zero
      - series: varnish_requests{instance="cp3052:3903", cluster="cache_text"}
        values: '0x11'
      - series: rdkafka_producer_topic_partition_msgs_total{instance="cp3052:9132", cluster="cache_text"}
        values: '0x11'
      # This should not fire an alert because the rate has been below 20% for at least 5 min
      - series: varnish_requests{instance="cp3053:3903", cluster="cache_text"}
        values: '0+20x8'
      - series: rdkafka_producer_topic_partition_msgs_total{instance="cp3053:9132", cluster="cache_text"}
        values: '0+1x7 180'
    external_labels:
      site: esams
    alert_rule_test:
    - alertname: VarnishkafkaNoMessages
      eval_time: 8m
      exp_alerts:
        - exp_labels:
            severity: warning
            team: data-engineering
            hostname: "cp3050"
            cluster: "cache_text"
          exp_annotations:
            summary: "[WARNING] - varnishkafka on cp3050 is not sending enough cache_text requests"
            description: "[WARNING] - Varnish on cp3050 is receiving requests, but Varnishkafka is not sending enough cache_text requests"
            dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?orgId=1&var-datasource=esams%20prometheus/ops&var-cp_cluster=cache_text&var-instance=cp3050%3A9132&viewPanel=14"
            runbook: &varnishkafka-nomessage-test-runbook "https://wikitech.wikimedia.org/wiki/Analytics/Systems/Varnishkafka"
        - exp_labels:
            severity: warning
            team: data-engineering
            hostname: "cp3051"
            cluster: "cache_text"
          exp_annotations:
            summary: "[WARNING] - varnishkafka on cp3051 is not sending enough cache_text requests"
            description: "[WARNING] - Varnish on cp3051 is receiving requests, but Varnishkafka is not sending enough cache_text requests"
            dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?orgId=1&var-datasource=esams%20prometheus/ops&var-cp_cluster=cache_text&var-instance=cp3051%3A9132&viewPanel=14"
            runbook: *varnishkafka-nomessage-test-runbook
        - exp_labels:
            severity: critical
            team: data-engineering
            hostname: "cp3051"
            cluster: "cache_text"
          exp_annotations:
            summary: "[CRITICAL] - varnishkafka on cp3051 is not sending enough cache_text requests"
            description: "[CRITICAL] - Varnish on cp3051 is receiving requests, but Varnishkafka is not sending enough cache_text requests"
            dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?orgId=1&var-datasource=esams%20prometheus/ops&var-cp_cluster=cache_text&var-instance=cp3051%3A9132&viewPanel=14"
            runbook: *varnishkafka-nomessage-test-runbook

  - interval: 1m
    external_labels:
      site: eqiad
    input_series:
      # Should not trigger as it averages fewer than 1 per second
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_upload", instance="cp3050:9132", source="webrequest"}'
        values: "0 30 60 90 120 150 180"
      # Same instance as above but a different cluster; Should not trigger even
      # though both combined would exceed warning threshold
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_text", instance="cp3050:9132", source="webrequest"}'
        values: "0 30 60 90 120 150 180"
      # Should not trigger by itself...
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_upload", instance="cp3051:9132", source="webrequest"}'
        values: "0 30 60 90 120 150 180"
      # ...but this one should add to the previous so it *does* trigger (we
      # don't care about separating by source)!
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_upload", instance="cp3051:9132", source="eventlogging"}'
        values: "0 30 60 90 120 150 180"
      # Should trigger a warning
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_text", instance="cp3052:9132", source="webrequest"}'
        values: "0 100 200 300 400 500 600"
      # Should trigger a separate alert since it's a different cluster
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_upload", instance="cp3052:9132", source="webrequest"}'
        values: "0 100 200 300 400 500 600"
      # More errors to trigger critical alert variant of the same alert name
      - series: 'varnishkafka_delivery_errors_total{cluster="cache_text", instance="cp3053:9132", source="webrequest"}'
        values: "0 500 1000 1500 2000 2500 3000"
    alert_rule_test:
      - alertname: VarnishKafkaDeliveryErrors
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              cluster: "cache_upload"
              instance: "cp3051:9132"
              severity: "warning"
              team: "data-engineering"
            exp_annotations:
              summary: "varnishkafka has cache_upload errors on cp3051:9132"
              description: "varnishkafka is experiencing cache_upload errors on the instance cp3051:9132 in the eqiad datacenter."
              dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?panelId=20&fullscreen&orgId=1&var-datasource=eqiad%20prometheus/ops&var-cp_cluster=cache_upload&var-instance=cp3051"
              runbook: &wt-varnishkafka "https://wikitech.wikimedia.org/wiki/Analytics/Systems/Varnishkafka"
          - exp_labels:
              cluster: "cache_text"
              instance: "cp3052:9132"
              severity: "warning"
              team: "data-engineering"
            exp_annotations:
              summary: "varnishkafka has cache_text errors on cp3052:9132"
              description: "varnishkafka is experiencing cache_text errors on the instance cp3052:9132 in the eqiad datacenter."
              dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?panelId=20&fullscreen&orgId=1&var-datasource=eqiad%20prometheus/ops&var-cp_cluster=cache_text&var-instance=cp3052"
              runbook: *wt-varnishkafka
          - exp_labels:
              cluster: "cache_upload"
              instance: "cp3052:9132"
              severity: "warning"
              team: "data-engineering"
            exp_annotations:
              summary: "varnishkafka has cache_upload errors on cp3052:9132"
              description: "varnishkafka is experiencing cache_upload errors on the instance cp3052:9132 in the eqiad datacenter."
              dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?panelId=20&fullscreen&orgId=1&var-datasource=eqiad%20prometheus/ops&var-cp_cluster=cache_upload&var-instance=cp3052"
              runbook: *wt-varnishkafka
          - exp_labels:
              cluster: "cache_text"
              instance: "cp3053:9132"
              severity: "warning"
              team: "data-engineering"
            exp_annotations: &varnishkafka-cp3053-annotation
              summary: "varnishkafka has cache_text errors on cp3053:9132"
              description: "varnishkafka is experiencing cache_text errors on the instance cp3053:9132 in the eqiad datacenter."
              dashboard: "https://grafana.wikimedia.org/d/000000253/varnishkafka?panelId=20&fullscreen&orgId=1&var-datasource=eqiad%20prometheus/ops&var-cp_cluster=cache_text&var-instance=cp3053"
              runbook: *wt-varnishkafka
          - exp_labels:
              cluster: "cache_text"
              instance: "cp3053:9132"
              severity: "critical"
              team: "data-engineering"
            exp_annotations: *varnishkafka-cp3053-annotation
