rule_files:
  - interfaces.yaml
evaluation_interval: 2m
tests:
  - interval: 2m
    external_labels:
      site: ulsfo
    input_series:
      - series: 'gnmi_interfaces_interface_state_oper_status{instance="cr3-ulsfo:9804", interface_name="xe-0/1/7", interface_description="Core: cr4-ulsfo:xe-0/1/7 {#1074}"}'
        values: '2 2'
      - series: 'gnmi_interfaces_interface_state_enabled{instance="cr3-ulsfo:9804", interface_name="xe-0/1/7", interface_description="Core: cr4-ulsfo:xe-0/1/7 {#1074}"}'
        values: '1 1'
      - series: 'gnmi_interfaces_interface_state_counters_in_errors{instance="cr3-ulsfo:9804", interface_name="xe-0/1/7", interface_description="Core: cr4-ulsfo:xe-0/1/7 {#1074}"}'
        values: '0+4x15'
      - series: 'gnmi_interfaces_interface_state_counters_in_fcs_errors{instance="cr3-ulsfo:9804", interface_name="xe-0/1/7", interface_description="Core: cr4-ulsfo:xe-0/1/7 {#1074}"}'
        values: '0+4x15'
      - series: 'gnmi_interfaces_interface_state_counters_out_errors{instance="cr3-ulsfo:9804", interface_name="xe-0/1/7", interface_description="Core: cr4-ulsfo:xe-0/1/7 {#1074}"}'
        values: '0+4x15'
      - series: 'gnmi_interfaces_interface_state_counters_out_octets{instance="cr2-esams:9804", interface_description="Peering: Test peering", interface_name="xe-0/1/0"}'
        values: '0+144000000000x5' # 1200000000 * 2 * 60
      - series: 'gnmi_interfaces_interface_state_counters_in_octets{instance="cr2-esams:9804", interface_description="Peering: Test peering", interface_name="xe-0/1/0"}'
        values: '0+144000000000x5' # 1200000000 * 2 * 60
      - series: 'gnmi_interfaces_interface_state_high_speed{instance="cr2-esams:9804", interface_description="Peering: Test peering", interface_name="xe-0/1/0"}'
        values: '10000x5'
      - series: 'gnmi_interfaces_interface_state_oper_status{instance="asw1-b12-drmrs:9804", interface_name="et-0/0/48", interface_description="Core: cr1-drmrs:et-0/0/1 {#D0100}"}'
        values: '2 2'
      - series: 'gnmi_interfaces_interface_state_enabled{instance="asw1-b12-drmrs:9804", interface_name="et-0/0/48", interface_description="Core: cr1-drmrs:et-0/0/1 {#D0100}"}'
        values: '1 1'
    alert_rule_test:
      - alertname: CoreRouterInterfaceDown
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              instance: cr3-ulsfo:9804
              interface_name: xe-0/1/7
              interface_description: "Core: cr4-ulsfo:xe-0/1/7 {#1074}"
              team: sre
              severity: critical
              scope: network
            exp_annotations:
              summary: "Core router interface down - cr3-ulsfo:xe-0/1/7 (Core: cr4-ulsfo:xe-0/1/7 {#1074})"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/fb403d62-5f03-434a-9dff-bd02b9fff504/network-device-overview?var-instance=cr3-ulsfo:9804'
              runbook: 'https://wikitech.wikimedia.org/wiki/Network_monitoring#Router_interface_down'
      - alertname: InboundInterfaceErrors
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              instance: cr3-ulsfo:9804
              interface_name: xe-0/1/7
              interface_description: "Core: cr4-ulsfo:xe-0/1/7 {#1074}"
              team: dcops
              severity: task
              scope: network
            exp_annotations:
              summary: "Inbound errors on interface cr3-ulsfo:xe-0/1/7 (Core: cr4-ulsfo:xe-0/1/7 {#1074})"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/fb403d62-5f03-434a-9dff-bd02b9fff504/network-device-overview?var-instance=cr3-ulsfo:9804'
              runbook: 'https://wikitech.wikimedia.org/wiki/Network_monitoring#Inbound/outbound_interface_errors'
      - alertname: OutboundInterfaceErrors
        eval_time: 20m
        exp_alerts:
          - exp_labels:
              instance: cr3-ulsfo:9804
              interface_name: xe-0/1/7
              interface_description: "Core: cr4-ulsfo:xe-0/1/7 {#1074}"
              team: dcops
              severity: task
              scope: network
            exp_annotations:
              summary: "Outbound errors on interface cr3-ulsfo:xe-0/1/7 (Core: cr4-ulsfo:xe-0/1/7 {#1074})"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/fb403d62-5f03-434a-9dff-bd02b9fff504/network-device-overview?var-instance=cr3-ulsfo:9804'
              runbook: 'https://wikitech.wikimedia.org/wiki/Network_monitoring#Inbound/outbound_interface_errors'
      - alertname: TransitPeeringOutboundSaturation
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              instance: cr2-esams:9804
              interface_name: xe-0/1/0
              interface_description: "Peering: Test peering"
              team: sre
              severity: page
              scope: network
            exp_annotations:
              summary: "Transit or peering outbound traffic above 90% capacity - cr2-esams:xe-0/1/0 (Peering: Test peering) #page"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/fb403d62-5f03-434a-9dff-bd02b9fff504/network-device-overview?var-instance=cr2-esams:9804'
              runbook: 'https://wikitech.wikimedia.org/wiki/Network_monitoring#Primary_outbound_port_utilization_over_90%'
      - alertname: TransitPeeringInboundSaturation
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              instance: cr2-esams:9804
              interface_name: xe-0/1/0
              interface_description: "Peering: Test peering"
              team: sre
              severity: page
              scope: network
            exp_annotations:
              summary: "Transit or peering inbound traffic above 90% capacity - cr2-esams:xe-0/1/0 (Peering: Test peering) #page"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/fb403d62-5f03-434a-9dff-bd02b9fff504/network-device-overview?var-instance=cr2-esams:9804'
              runbook: 'https://wikitech.wikimedia.org/wiki/Network_monitoring#Primary_inbound_port_utilization_over_90%'
      - alertname: SwitchCoreInterfaceDown
        eval_time: 2m
        exp_alerts:
          - exp_labels:
              instance: asw1-b12-drmrs:9804
              interface_name: et-0/0/48
              interface_description: "Core: cr1-drmrs:et-0/0/1 {#D0100}"
              team: sre
              severity: critical
              scope: network
            exp_annotations:
              summary: "Switch core interface down - asw1-b12-drmrs:et-0/0/48 (Core: cr1-drmrs:et-0/0/1 {#D0100})"
              description: ""
              dashboard: 'https://grafana.wikimedia.org/d/fb403d62-5f03-434a-9dff-bd02b9fff504/network-device-overview?var-instance=asw1-b12-drmrs:9804'
              runbook: 'https://wikitech.wikimedia.org/wiki/Network_monitoring#Switch_interface_down'