rule_files:
  - vrts.yaml
evaluation_interval: 5m

tests:
  # HighInodeUsage
  - interval: 1m
    input_series:
      - series: 'node_filesystem_files_free{instance="vrts1003:9100", fstype="ext4"}'
        values: '500000'
      - series: 'node_filesystem_files{instance="vrts1003:9100", fstype="ext4"}'
        values: '4890624'
    alert_rule_test:
      - alertname: VRTSHighInodeUsage
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: collaboration-services
              instance: vrts1003:9100
              fstype: ext4
            exp_annotations:
              summary: High inode usage on vrts1003:9100
              description: vrts1003:9100 inodes are almost exhausted
              runbook: https://wikitech.wikimedia.org/wiki/VRT_System/Troubleshooting#VRTSHighInodeUsage
              dashboard: https://grafana.wikimedia.org/goto/jbpjcxZvg?orgId=1
  - interval: 1m
    input_series:
      - series: 'sql_vrts_ticket_count{name="junk_queues"}'
        values: '30000'
    alert_rule_test:
      - alertname: VRTSJunkQueueSize
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              severity: warning
              team: collaboration-services
              name: junk_queues
            exp_annotations:
              summary: VRTS junk queue size is large
              description: The VRTS junk queue is above its usual threshold
              runbook: https://wikitech.wikimedia.org/wiki/VRT_System/Troubleshooting#VRTSJunkQueueSize
              dashboard: https://grafana.wikimedia.org/goto/jbpjcxZvg?orgId=1