rule_files:
  - haproxykafka.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    external_labels:
      site: eqsin
    input_series:
      - series: 'up{instance="cp5017:9341"}'
        values: '1 1 1 1 0 0 0 0 0 0 0'  # Goes down at 4m, stays down
    alert_rule_test:
      - eval_time: 4m
        alertname: HaproxyKafkaExporterDown
        exp_alerts: []  # No alert yet
      - eval_time: 9m
        alertname: HaproxyKafkaExporterDown
        exp_alerts:
         - exp_labels:
             severity: warning
             team: traffic
             instance: "cp5017:9341"
           exp_annotations:
             summary: 'HaproxyKafka on cp5017 is down'
             description: 'HaproxyKafka on cp5017 has been down for more than 5 minutes'
             dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
             runbook: "https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaExporterDown"

  - interval: 1m
    external_labels:
      site: eqsin
    input_series:
      - series: 'up{instance="cp5017:9341"}'
        values: '1 1 1 0+0x70' # 70m
    alert_rule_test:
      - eval_time: 1h3m
        alertname: HaproxyKafkaExporterDown
        exp_alerts:
         - exp_labels:
             severity: critical
             team: traffic
             instance: "cp5017:9341"
           exp_annotations:
             summary: 'HaproxyKafka on cp5017 is down'
             description: 'HaproxyKafka on cp5017 has been down for more than 1 hour'
             dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
             runbook: "https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaExporterDown"
         - exp_labels:
             severity: warning
             team: traffic
             instance: "cp5017:9341"
           exp_annotations:
             summary: 'HaproxyKafka on cp5017 is down'
             description: 'HaproxyKafka on cp5017 has been down for more than 5 minutes'
             dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
             runbook: "https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaExporterDown"

  - interval: 1m
    external_labels:
      site: eqsin
    # test case 1, WARNING when ratio < 0.2 and >= 0.1
    input_series:
      # sum by (hostname, cluster) (label_replace (rate(haproxykafka_valid_messages_total[5m]), "hostname", "$1", "instance", "(.*):.*"))
      # /
      # sum by (hostname, cluster) (label_replace (rate(haproxy_frontend_http_requests_total[5m]), "hostname", "$1", "instance", "(.*):.*")) < 0.2
      - series: 'haproxykafka_valid_messages_total{instance="cp5017:9341", cluster="cache_text", site="eqsin"}'
        values: '0+2x5'  # Rate of 10/5m = 2 per minute
      - series: 'haproxy_frontend_http_requests_total{instance="cp5017:9422", cluster="cache_text", site="eqsin"}'
        values: '0+20x5' # Rate of 100/5m = 20 per minute
    alert_rule_test:
      - alertname: HaproxyKafkaNoMessages
        eval_time: 6m
        exp_alerts:
         - exp_labels:
             hostname: "cp5017"
             severity: warning
             team: traffic
             cluster: cache_text
           exp_annotations:
              summary: 'Unexpected rate of produced HaproxyKafka messages by cp5017'
              description: 'Haproxy on cp5017 is receiving requests, but HaproxyKafka is not sending enough messages'
              dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
              runbook: 'https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaNoMessages'

  - interval: 1m
    external_labels:
      site: eqsin
    # Test case 2: CRITICAL when ratio < 0.1
    input_series:
      - series: 'haproxykafka_valid_messages_total{instance="cp5017:9341", cluster="cache_text", site="eqsin"}'
        values: '0+1x10'   # Rate of 1/5m = 0.2 per minute
      - series: 'haproxy_frontend_http_requests_total{instance="cp5017:9422", cluster="cache_text", site="eqsin"}'
        values: '0+100x10' # Rate of 100/5m = 20 per minute
    alert_rule_test:
      # Test case 2: CRITICAL alert should fire when ratio is < 0.1
      - eval_time: 6m
        alertname: HaproxyKafkaNoMessages
        exp_alerts:
          - exp_labels:
              hostname: "cp5017"
              severity: warning
              team: traffic
              cluster: cache_text
            exp_annotations:
              summary: 'Unexpected rate of produced HaproxyKafka messages by cp5017'
              description: 'Haproxy on cp5017 is receiving requests, but HaproxyKafka is not sending enough messages'
              dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
              runbook: 'https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaNoMessages'
          - exp_labels:
              severity: critical
              team: traffic
              hostname: "cp5017"
              cluster: cache_text
            exp_annotations:
              summary: 'Unexpected rate of produced HaproxyKafka messages by cp5017'
              description: 'Haproxy on cp5017 is receiving requests, but HaproxyKafka is not sending enough messages'
              dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
              runbook: 'https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaNoMessages'
  - interval: 1m
    external_labels:
      site: eqsin
    # Test case 3: haproxykafka_valid_messages_total metric is missing
    input_series:
      - series: 'up{instance="cp5017:9341", cluster="cache_text", site="eqsin"}'
        values: '0x10'
      - series: 'haproxy_frontend_http_requests_total{instance="cp5017:9422", cluster="cache_text", site="eqsin"}'
        values: '0+100x10' # Rate of 100/5m = 20 per minute
    alert_rule_test:
      # Test case 3: CRITICAL alert should fire when ratio is < 0.1
      - eval_time: 6m
        alertname: HaproxyKafkaNoMessages
        exp_alerts:
          - exp_labels:
              hostname: "cp5017"
              severity: warning
              team: traffic
              cluster: cache_text
            exp_annotations:
              summary: 'Unexpected rate of produced HaproxyKafka messages by cp5017'
              description: 'Haproxy on cp5017 is receiving requests, but HaproxyKafka is not sending enough messages'
              dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
              runbook: 'https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaNoMessages'
          - exp_labels:
              severity: critical
              team: traffic
              hostname: "cp5017"
              cluster: cache_text
            exp_annotations:
              summary: 'Unexpected rate of produced HaproxyKafka messages by cp5017'
              description: 'Haproxy on cp5017 is receiving requests, but HaproxyKafka is not sending enough messages'
              dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
              runbook: 'https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaNoMessages'

  - interval: 1m
    external_labels:
      site: eqsin
    input_series:
      - series: 'node_systemd_service_restart_total{instance="cp5017:9100", cluster="cache_text", name="haproxykafka.service"}'
        values: '0 1 2 3 4 5 6 7'
      - series: 'node_systemd_service_restart_total{instance="cp5018:9100", cluster="cache_text", name="haproxykafka.service"}'
        values: '0 1 0 0 0 0 0 0'
    alert_rule_test:
      - alertname: HaproxyKafkaRestarted
        eval_time: 7m
        exp_alerts:
          - exp_labels:
              severity: critical
              team: traffic
              instance: cp5017:9100
              cluster: cache_text
              name: haproxykafka.service
            exp_annotations:
              summary: 'HaproxyKafka restarted on cp5017:9100'
              description: "HaproxyKafka on cp5017:9100 (eqsin) has restarted unexpectedly."
              dashboard: 'https://grafana.wikimedia.org/d/d3e4e37c-c1d9-47af-9aad-a08dae2b3fd5/haproxykafka?orgId=1&var-site=eqsin&var-instance=cp5017'
              runbook: 'https://wikitech.wikimedia.org/wiki/HAProxyKafka#HaproxyKafkaRestarted'
