rule_files:
  - varnish.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'up{instance="cp3030:9331", job="varnish-text"}'
        values: '1 0 0 0 0 0 0'
    alert_rule_test:
    - alertname: VarnishPrometheusExporterDown
      eval_time: 6m
      exp_alerts:
       - exp_labels:
           severity: warning
           team: traffic
           instance: "cp3030:9331"
           job: "varnish-text"
         exp_annotations:
           summary: "Varnish Exporter on instance cp3030:9331 is unreachable"
           description: "Prometheus has been unable to fetch metrics from Varnish Exporter on host cp3030:9331 job(varnish-text) for more than 5 minutes. Make sure the exporter is running on the host."
           dashboard: 'https://grafana.wikimedia.org/d/000000304/varnish-dc-stats?viewPanel=17'
           runbook: 'https://wikitech.wikimedia.org/wiki/Prometheus#Prometheus_job_unavailable'

  - interval: 1m
    external_labels:
      site: eqiad
    input_series:
      # Should trigger an alert
      - series: 'varnish_mgt_child_start{instance="cp1234:9331", job="varnish-text", layer="frontend"}'
        values: "1 2 1 1 1 1 1"
      # Should not trigger an alert
      - series: 'varnish_mgt_child_start{instance="cp1235:9331", job="varnish-upload", layer="frontend"}'
        values: "1 1 1 1 1 1 1"
      # Should also not trigger an alert; This data would suggest that the
      # service is disabled (and monitoring this would be for another alert)
      - series: 'varnish_mgt_child_start{instance="cp1236:9331", job="varnish-upload", layer="frontend"}'
        values: "0 0 0 0 0 0 0"
    alert_rule_test:
      - alertname: VarnishChildRestarted
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              instance: "cp1234:9331"
              job: "varnish-text"
              layer: "frontend"
              severity: "critical"
              team: "traffic"
            exp_annotations:
              summary: "varnish-text restarted on cp1234"
              description: "Varnish seems to be unstable on cp1234, specifically varnish-text."
              dashboard: "https://grafana.wikimedia.org/d/000000330/varnish-machine-stats?orgId=1&viewPanel=66&var-server=cp1234&datasource=eqiad%20prometheus/ops"
              runbook: "https://wikitech.wikimedia.org/wiki/Varnish"

  - interval: 1m
    external_labels:
      site: eqiad
    input_series:
      # Not quite enough for an alert but close
      - series: 'varnishd_mmap_count{instance="cp1234:9100", cluster="cache_upload", site="eqiad"}'
        values: '4000 5000 4000 5000 4000 5000 4000'
      - series: 'sysctl_vm_max_map_count{instance="cp1234:9100", cluster="cache_upload", site="eqiad"}'
        values: '6000 6000 6000 6000 6000 6000 6000'
      # But this one is!
      - series: 'varnishd_mmap_count{instance="cp1235:9100", cluster="cache_upload", site="eqiad"}'
        values: '4000 5000 5000 5000 5000 5000 5000'
      - series: 'sysctl_vm_max_map_count{instance="cp1235:9100", cluster="cache_upload", site="eqiad"}'
        values: '6000 6000 6000 6000 6000 6000 6000'
      # Critical limits
      - series: 'varnishd_mmap_count{instance="cp1236:9100", cluster="cache_upload", site="eqiad"}'
        values: '5000 5500 5500 5500 5500 5500 5500'
      - series: 'sysctl_vm_max_map_count{instance="cp1236:9100", cluster="cache_upload", site="eqiad"}'
        values: '6000 6000 6000 6000 6000 6000 6000'
    alert_rule_test:
      - alertname: VarnishHighMmapCount
        eval_time: 6m
        exp_alerts:
         - exp_labels:
             severity: warning
             team: traffic
             instance: "cp1235:9100"
             cluster: "cache_upload"
             site: eqiad
           exp_annotations:
             summary: "Varnish's number of memory map areas are reaching the maximum limit"
             description: "Varnish on cp1235:9100 (eqiad) is at 83% of the vm.max_map_count limit."
             dashboard: 'https://grafana.wikimedia.org/d/wiU3SdEWk/cache-host-drilldown?viewPanel=76&var-site=eqiad&var-instance=cp1235'
             runbook: "https://wikitech.wikimedia.org/wiki/Varnish"
         - exp_labels:
             severity: warning
             team: traffic
             instance: "cp1236:9100"
             cluster: "cache_upload"
             site: eqiad
           exp_annotations:
             summary: "Varnish's number of memory map areas are reaching the maximum limit"
             description: "Varnish on cp1236:9100 (eqiad) is at 92% of the vm.max_map_count limit."
             dashboard: 'https://grafana.wikimedia.org/d/wiU3SdEWk/cache-host-drilldown?viewPanel=76&var-site=eqiad&var-instance=cp1236'
             runbook: "https://wikitech.wikimedia.org/wiki/Varnish"
         - exp_labels:
             severity: critical
             team: traffic
             instance: "cp1236:9100"
             cluster: "cache_upload"
             site: eqiad
           exp_annotations:
             summary: "Varnish's number of memory map areas are reaching the maximum limit"
             description: "Varnish on cp1236:9100 (eqiad) is at 92% of the vm.max_map_count limit."
             dashboard: 'https://grafana.wikimedia.org/d/wiU3SdEWk/cache-host-drilldown?viewPanel=76&var-site=eqiad&var-instance=cp1236'
             runbook: "https://wikitech.wikimedia.org/wiki/Varnish"

  - interval: 1m
    external_labels:
      site: eqiad
    input_series:
      - series: 'varnish_main_threads_used_percent{instance="cp1234:0"}'
        values: '60 69'
      - series: 'varnish_main_threads_used_percent{instance="cp1235:0"}'
        values: '60 70'
      - series: 'varnish_main_threads_used_percent{instance="cp1236:0"}'
        values: '70 80'
    alert_rule_test:
      - alertname: VarnishHighThreadCount
        eval_time: 2m
        exp_alerts:
         - exp_labels:
             severity: warning
             team: traffic
             instance: "cp1235:0"
           exp_annotations:
             summary: "Varnish's thread count on cp1235:0 is high"
             description: "Varnish on cp1235:0 (eqiad) is at 70% of the maximum limit."
             dashboard: 'https://grafana.wikimedia.org/d/wiU3SdEWk/cache-host-drilldown?viewPanel=99&var-site=eqiad&var-instance=cp1235'
             runbook: "https://wikitech.wikimedia.org/wiki/Varnish"
         - exp_labels:
             severity: warning
             team: traffic
             instance: "cp1236:0"
           exp_annotations:
             summary: "Varnish's thread count on cp1236:0 is high"
             description: "Varnish on cp1236:0 (eqiad) is at 80% of the maximum limit."
             dashboard: 'https://grafana.wikimedia.org/d/wiU3SdEWk/cache-host-drilldown?viewPanel=99&var-site=eqiad&var-instance=cp1236'
             runbook: "https://wikitech.wikimedia.org/wiki/Varnish"
         - exp_labels:
             severity: critical
             team: traffic
             instance: "cp1236:0"
           exp_annotations:
             summary: "Varnish's thread count on cp1236:0 is high"
             description: "Varnish on cp1236:0 (eqiad) is at 80% of the maximum limit."
             dashboard: 'https://grafana.wikimedia.org/d/wiU3SdEWk/cache-host-drilldown?viewPanel=99&var-site=eqiad&var-instance=cp1236'
             runbook: "https://wikitech.wikimedia.org/wiki/Varnish"
